```scala

```

# CSCI 3155 Homework 3

More practice with algebraic data types

__Name__:

Qiuyang Fu

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (40 Points) : Useful functions on natural numbers

The type Nat we have defined allows us to create objects that are isomorphic to the natural numbers. However we're still missing many functions that allow us to use them like we do regular numbers.

```scala
sealed trait Nat
case object Zero extends Nat
case class Succ(pred : Nat) extends Nat

val one   = Succ(Zero)
val two   = Succ(one)
val three = Succ(two)
val four  = Succ(three)
val five  = Succ(four)
val six   = Succ(five)
val seven = Succ(six)
val eight = Succ(seven)
val nine  = Succ(eight)
val ten   = Succ(nine)
```

Output:
```
defined [32mtrait[39m [36mNat[39m
defined [32mobject[39m [36mZero[39m
defined [32mclass[39m [36mSucc[39m
[36mone[39m: [32mSucc[39m = [33mSucc[39m(Zero)
[36mtwo[39m: [32mSucc[39m = [33mSucc[39m([33mSucc[39m(Zero))
[36mthree[39m: [32mSucc[39m = [33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero)))
[36mfour[39m: [32mSucc[39m = [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))
[36mfive[39m: [32mSucc[39m = [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero)))))
[36msix[39m: [32mSucc[39m = [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))))
[36mseven[39m: [32mSucc[39m = [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero)))))))
[36meight[39m: [32mSucc[39m = [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))))))
[3
```

### 1a - conversion
Write three conversion functions, ``nat2int``, ``int2nat``, and ``nat2str``.

``nat2int`` : Nat $\rightarrow$ Int.  This function takes a Nat and returns the integer it is equal to.

``int2nat`` : Int $\rightarrow$ Nat.  This function takes a integer and returns the nat it is equal to. What should you do for integers that don't have a corresponding Nat?

``nat2str`` : Nat $\rightarrow$ String. Hint integers have a handy .toString method

<!-- grade_id: cell-59a9eb7a40472fe3, points: 0, grade: false, solution: true -->
```scala
def nat2int(x : Nat) : Int = 
    x match {
        case Zero => 0
        case Succ(p) => nat2int(p) + 1
    }

def nat2str(x : Nat) : String = nat2int(x).toString

def int2nat(x : Int) : Nat = 
    x match {
        case 0 => Zero
        case _ => Succ(int2nat(x - 1))
    }
```

Output:
```
defined [32mfunction[39m [36mnat2int[39m
defined [32mfunction[39m [36mnat2str[39m
defined [32mfunction[39m [36mint2nat[39m
```

<!-- grade_id: cell-53901f335b3392b0, points: 10, grade: true, solution: false -->
```scala
assert(nat2int(two) == 2)
assert(int2nat(nat2int(Zero)) ==  Zero)
assert(int2nat(nat2int(five)) == five)
assert(nat2str(Zero) == "0")
assert(nat2str(five) == "5")
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
```

### 1b - checking equality
Write ``eq_nat`` : (Nat, Nat) => Boolean, which determines whether two nats are equal. Remember you can match on a pair of arguments if your cases are also pairs.

<!-- grade_id: cell-bc37e646de510c68, points: 0, grade: false, solution: true -->
```scala
def eq_nat(x : Nat, y : Nat) : Boolean = x == y
```

Output:
```
defined [32mfunction[39m [36meq_nat[39m
```

<!-- grade_id: cell-e8860a2cfca1a9f7, points: 10, grade: true, solution: false -->
```scala
assert( eq_nat(Zero,Zero))
assert( eq_nat(five,five))
assert(!eq_nat(Zero,five))
assert(!eq_nat(one,Zero))
assert(!eq_nat(four,five))
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
```

### 1c - $\neq$
Write ``neq_nat``, which returns false if both argurments are equal and true if they are not equal.

<!-- grade_id: cell-34671b5db948a24d, points: 0, grade: false, solution: true -->
```scala
def neq_nat(x : Nat, y : Nat) : Boolean = x != y
```

Output:
```
defined [32mfunction[39m [36mneq_nat[39m
```

<!-- grade_id: cell-7f9d5faa51b1ca57, points: 5, grade: true, solution: false -->
```scala
assert(!neq_nat(Zero,Zero))
assert(!neq_nat(five,five))
assert( neq_nat(Zero,five))
assert( neq_nat(one,Zero))
assert( neq_nat(four,five))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

### 1d - $\le$
Write a less than or equal to function for nats.

$x \le y$  : : ``lte_nat(x,y)``

<!-- grade_id: cell-3bbd75c47b274054, points: 0, grade: false, solution: true -->
```scala
def lte_nat(x : Nat, y : Nat) : Boolean = nat2int(x) <= nat2int(y)
```

Output:
```
defined [32mfunction[39m [36mlte_nat[39m
```

<!-- grade_id: cell-780a49d05042a4b9, points: 10, grade: true, solution: false -->
```scala
assert(lte_nat(Zero,Zero))
assert(lte_nat(Zero,one))
assert(!lte_nat(one,Zero))
assert(lte_nat(four,four))
assert(lte_nat(four,five))
assert(!lte_nat(five,four))
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
```

### 1e - comparisons finished

Write $>$ ``gt_nat``, $<$ ``lt_nat``, and $\ge$ ``gte_nat``

<!-- grade_id: cell-f972d0fed798123f, points: 0, grade: false, solution: true -->
```scala
def gt_nat(x : Nat, y : Nat) : Boolean = nat2int(x) > nat2int(y)

def lt_nat(x : Nat, y : Nat) : Boolean = nat2int(x) < nat2int(y)

def gte_nat(x : Nat, y : Nat) : Boolean = nat2int(x) >= nat2int(y)

```

Output:
```
defined [32mfunction[39m [36mgt_nat[39m
defined [32mfunction[39m [36mlt_nat[39m
defined [32mfunction[39m [36mgte_nat[39m
```

<!-- grade_id: cell-79c0e147c9def1bc, points: 5, grade: true, solution: false -->
```scala
assert( gt_nat(two,Zero))
assert(!gt_nat(two,three))
assert(!gt_nat(three,three))
assert(!lt_nat(two,Zero))
assert( lt_nat(two,three))
assert(!lt_nat(three,three))
assert(!lte_nat(two,Zero))
assert( lte_nat(two,three))
assert( lte_nat(three,three))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

## Problem 2 (15 points) : Math on Nats

```scala
def plus(x : Nat, y : Nat) : Nat =
    x match {
        case Zero    => y
        case Succ(p) => Succ(plus(p, y))
    }

plus(two,three) == five

def mult(x : Nat, y : Nat) : Nat = 
    x match {
        case Zero    => Zero
        case Succ(p) => plus(mult(p,y), y)
    }

mult(two, four) == plus(four,four)
```

Output:
```
defined [32mfunction[39m [36mplus[39m
[36mres12_1[39m: [32mBoolean[39m = true
defined [32mfunction[39m [36mmult[39m
[36mres12_3[39m: [32mBoolean[39m = true
```

### 2a - power

Define a ``pow`` function that raise a nat to a power.

$x ^ y $ = ``pow(x,y)``

Hint: you can use function you've already written

$$pow : (\mathbb{N},\mathbb{N}) \rightarrow \mathbb{N}$$

<!-- grade_id: cell-2355e38d2b3ce4e2, points: 0, grade: false, solution: true -->
```scala
def pow(x : Nat, y : Nat) : Nat = 
    y match {
         case Zero => one
         case _ => mult(x, pow(x, int2nat(nat2int(y)-1)))
    }
```

Output:
```
defined [32mfunction[39m [36mpow[39m
```

<!-- grade_id: cell-1ec0112c26ef161b, points: 5, grade: true, solution: false -->
```scala
assert(pow(three, two) == nine)
assert(pow(four, Zero) == one)
assert(pow(two, two) == four)
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

### 2b - minus on Nats

Write the minus function, which subtracts a natural number from another natural number. Note that we cannot have negative numbers as a Nat. If the result would be less than zero then leave the answer as zero.


$$minus : (\mathbb{N},\mathbb{N}) \rightarrow \mathbb{N}$$

<!-- grade_id: cell-c07e902c2030a5c0, points: 0, grade: false, solution: true -->
```scala
def minus(x : Nat, y : Nat) : Nat = {
    var cntx = nat2int(x)
    var cnty = nat2int(y)
    if(cntx < cnty) {
        return Zero
    }
    return int2nat(cntx - cnty)
}
```

Output:
```
defined [32mfunction[39m [36mminus[39m
```

<!-- grade_id: cell-f9ef1004d9f8a50e, points: 10, grade: true, solution: false -->
```scala
assert(minus(four, two) == two)
assert(minus(two, four) == Zero)
assert(minus(four, four) == Zero)
assert(minus(four, Zero) == four)
assert(minus(Zero, four) == Zero)
assert(minus(Zero,Zero) == Zero)
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
```

## Problem 3 (25 points) : Lists

```scala
sealed trait NatList
case object Empty extends NatList
case class Cons(h : Nat, t : NatList) extends NatList

val list1 = Empty
val list2 = Cons(five, Empty)
val list3 = Cons(three, Cons(one, Cons(four, Cons(one, Cons(five, Empty)))))
```

Output:
```
defined [32mtrait[39m [36mNatList[39m
defined [32mobject[39m [36mEmpty[39m
defined [32mclass[39m [36mCons[39m
[36mlist1[39m: [32mEmpty[39m.type = Empty
[36mlist2[39m: [32mCons[39m = [33mCons[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))), Empty)
[36mlist3[39m: [32mCons[39m = [33mCons[39m(
  [33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))),
  [33mCons[39m(
    [33mSucc[39m(Zero),
    [33mCons[39m(
      [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero)))),
      [33mCons[39m([33mSucc[39m(Zero), [33mCons[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))), Empty))
    )
  )
)
```

### 3b - prod

write a ``prod`` function which takes the product of all the elemts of a list as a Nat. Hint: an empty list should have prod = 1.


$$prod : \text{List } \mathbb{N} \rightarrow \mathbb{N}$$

<!-- grade_id: cell-e98c3f5a8ade1efb, points: 0, grade: false, solution: true -->
```scala
import scala.language.postfixOps
def prod(l : NatList) : Nat = 
    l match {
        case Empty => one
        case Cons(h, t) => mult(h, prod(t))
    }
```

Output:
```
[32mimport [39m[36mscala.language.postfixOps
[39m
defined [32mfunction[39m [36mprod[39m
```

<!-- grade_id: cell-58b209724daf6f0d, points: 5, grade: true, solution: false -->
```scala
assert(prod(list1) == one)
assert(prod(list2) == five)
assert(prod(list3) == mult(ten,six))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

### 3b - pretty

Write a function ``pretty`` that pretty prints a list. The function should return a string with the elements written with digits and separated by a single comma without spaces, and the string "[]" for the empty list

The ``list1`` above should print ``[]``

The ``list2`` above should print ``5``

The ``list3`` above should print ``3,1,4,1,5``

$$pretty : \text{List } \mathbb{N} \rightarrow String$$

<!-- grade_id: cell-2146893bbd87b1c8, points: 0, grade: false, solution: true -->
```scala
def pretty(l : NatList) : String = 
    l match {
        case Empty => "[]"
        case Cons(h, Empty) => nat2str(h)
        case Cons(h, t:NatList) => nat2str(h) + ',' + pretty(t)
    }
```

Output:
```
defined [32mfunction[39m [36mpretty[39m
```

<!-- grade_id: cell-2e4efbfb3b333cb2, points: 10, grade: true, solution: false -->
```scala
assert(pretty(list1) == "[]")
assert(pretty(list2) == "5")
assert(pretty(list3) == "3,1,4,1,5")
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
```

### 3c - contains

Write a function ``contains`` that returns true if a list contains an Nat, and false otherwise.

$$contains : \text{List } \mathbb{N} \rightarrow \mathbb{B}$$

<!-- grade_id: cell-f4e2053b7cc00cb9, points: 0, grade: false, solution: true -->
```scala
def contains(l : NatList, p : Nat) : Boolean = 
    l match {
        case Empty => false
        case Cons(h, t) if (h == p)  => true
        case Cons(h, t) if (h != p)  => contains(t, p)
    }
```

Output:
```
defined [32mfunction[39m [36mcontains[39m
```

<!-- grade_id: cell-97a594833affe029, points: 10, grade: true, solution: false -->
```scala
assert(!contains(list1, Zero))
assert(!contains(list2, Zero))
assert( contains(list2, five))
assert(!contains(list3, two))
assert( contains(list3, five))
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
```

```scala

```

```scala

```