```scala

```

# CSCI 3155 : Assignment 8

__Name__:

Qiuyang Fu

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 : Type Checking (100 points)

Write a function `typecheck` that typechecks a Lettuce program.

$$\texttt{typecheck} : Program \rightarrow Type$$ 

This function should take as input a lettuce _Program_, and return what _Type_ it is. If it doesn't typecheck it should return a special value _TypeError_. Use the Lettuce typing notes on moodle for details on how typechecking lettuce works.

The tests for this assigment are hidden, so you will need to write your own tests! Note that you are encouraged to collaborate and share your test ideas on Piazza. To get full points, your implementation needs to match the specification from the Lettuce typing notes on moodle.

```scala
sealed trait Type
case object NumType extends Type
case object BoolType extends Type
case class FunType( ty_a : Type, ty_r : Type) extends Type
case object TypeError extends Type

sealed trait Expr
case class  Const(n : Int) extends Expr
case class  Ident(s : String) extends Expr
case object True extends Expr
case object False extends Expr
case class  Plus( e1 : Expr, e2 : Expr) extends Expr
case class  Minus( e1 : Expr, e2 : Expr) extends Expr
case class  Mult( e1 : Expr, e2 : Expr) extends Expr
case class  Neg( e1 : Expr) extends Expr
case class  And( e1 : Expr, e2 : Expr) extends Expr
case class  Or( e1 : Expr, e2 : Expr) extends Expr
case class  Geq( e1 : Expr, e2 : Expr) extends Expr
case class  Eq( e1 : Expr, e2 : Expr) extends Expr
case class  ITE( ec : Expr, et : Expr, ef : Expr) extends Expr
case class  Let( name : String, ty_d : Type, ed : Expr, ei : Expr) extends Expr
case class  FunDef( arg : String, ty_a : Type, eb : Expr) extends Expr
case class  FunCall( ef : Expr, ea : Expr) extends Expr
case class  LetRec( f : String, ty_f : Type, x : String, ty_x : Type, eb : Expr, ei : Expr) extends Expr

sealed trait Program
case class Main(p : Expr) extends Program
```

Output:
```
defined [32mtrait[39m [36mType[39m
defined [32mobject[39m [36mNumType[39m
defined [32mobject[39m [36mBoolType[39m
defined [32mclass[39m [36mFunType[39m
defined [32mobject[39m [36mTypeError[39m
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mIdent[39m
defined [32mobject[39m [36mTrue[39m
defined [32mobject[39m [36mFalse[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mMult[39m
defined [32mclass[39m [36mNeg[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mGeq[39m
defined [32mclass[39m [36mEq[39m
defined [32mclass[39m [36mITE[39m
defined [32mclass[39m [36mLet[39m
defined [32mclass[39m [36mFunDef[39m
defined [32mclass[39m [36mFunCall[39m
defined [32mclass[39m [36mLetRec[39m
defined [32mtrait[39m [36mProgram[39m
defined [32mclass[39m [36mMain[39m
```

<!-- grade_id: cell-f8cd57266fd56cf2, points: 0, grade: false, solution: true -->
```scala
def typecheck(p : Program) = p match {
    case Main(e: Expr) => typecheckExpr(e, Map())
    case _ => TypeError
}

def typecheckExpr(e: Expr, env: Map[String, Type]) : Type = {
    def typeEquals(e1: Expr, t1: Type, e2: Expr, t2: Type, resultType: Type) = {
        val e1Type = typecheckExpr(e1, env)
        val e2Type = typecheckExpr(e2, env)
        if (e1Type == t1 && e2Type == t2) {
            resultType
        } else {
            TypeError
        }
    }
    e match {
        case Const(_)  => NumType
        case True => BoolType
        case False => BoolType
        case Ident(x: String) => env.getOrElse(x, TypeError)
        case Plus(e1: Expr, e2: Expr) => typeEquals(e1, NumType, e2, NumType, NumType)
        case Minus(e1: Expr, e2: Expr) => typeEquals(e1, NumType, e2, NumType, NumType)
        case Mult(e1: Expr, e2: Expr) => typeEquals(e1, NumType, e2, NumType, NumType)
        case And(e1: Expr, e2: Expr) => typeEquals(e1, BoolType, e2, BoolType, BoolType)
        case Or(e1: Expr, e2: Expr) => typeEquals(e1, BoolType, e2, BoolType, BoolType)
        case Neg(e: Expr) => typecheckExpr(e, env) match {
            case BoolType => BoolType
            case NumType => NumType
            case _ => TypeError
        }
        case Geq(e1: Expr, e2: Expr) => typeEquals(e1, NumType, e2, NumType, BoolType)
        case Eq(e1: Expr, e2: Expr) => (typecheckExpr(e1, env), typecheckExpr(e2, env)) match {
            case (TypeError, _) => TypeError
            case (_, TypeError) => TypeError
            case _ => BoolType
        }
        case ITE( ec : Expr, et : Expr, ef : Expr) => (typecheckExpr(ec, env), typecheckExpr(et, env), typecheckExpr(ef, env)) match {
            case (_, _, TypeError) => TypeError
            case (_, TypeError, _) => TypeError
            case (BoolType, t1, t2) => {
                if (t1 == t2) {
                    t1
                } else {
                    TypeError
                }
            }
            case _ => TypeError
        }
        case Let( name : String, ty_d : Type, ed : Expr, ei : Expr) => {
            if (typecheckExpr(ed, env) == ty_d){
                val new_env = env + (name -> ty_d)
                typecheckExpr(ei, new_env)
            } else {
                TypeError
            }
        }
        case FunDef( arg : String, ty_a : Type, eb : Expr) => {
            val new_env = env + (arg -> ty_a)
            val new_t = typecheckExpr(eb, new_env)
            if (new_t == TypeError) {
                TypeError
            } else {
                FunType(ty_a, new_t)
            }
        }
        case FunCall( ef : Expr, ea : Expr) =>  typecheckExpr(ef, env) match {
            case FunType(ty_a, ty_r) => {
                if (typecheckExpr(ea, env) == ty_a){
                    ty_r
                } else {
                    TypeError
                }
            }
            case _ => TypeError
        }
        case LetRec( f : String, ty_f : Type, x : String, ty_x : Type, eb : Expr, ei : Expr) => {
            val new_env = env + (x -> ty_f, f -> FunType(ty_f, ty_x))
            val new_t = typecheckExpr(eb, new_env)
            if (new_t == ty_x){
                typecheckExpr(ei, env + (f -> FunType(ty_f, ty_x)))
            } else {
                TypeError
            }
        }
        case _ => TypeError
    }
}
```

Output:
```
defined [32mfunction[39m [36mtypecheck[39m
defined [32mfunction[39m [36mtypecheckExpr[39m
```

<!-- grade_id: cell-d9b206afd7a24388, points: 20, grade: true, solution: false -->
```scala
val p1 : Program = Main( Const(10) )
assert( typecheck(p1) == NumType, "test 1")
passed(20)
println("Hidden Tests : up to 80 points remaining")
```

Output:
```
*** Tests Passed (20 points) ***
Hidden Tests : up to 80 points remaining
[36mp1[39m: [32mProgram[39m = [33mMain[39m([33mConst[39m([32m10[39m))
```

```scala
// test Const 
val p1 = Main( True )
val p2 = Main( False )
assert( typecheck(p1) == BoolType, "test 1")
assert( typecheck(p2) == BoolType, "test 2")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m(True)
[36mp2[39m: [32mMain[39m = [33mMain[39m(False)
```

```scala
// test Plus
val p1 = Main( Plus(Const(1), Const(2)) )
val p2 = Main( Plus(Const(1), False) )
assert( typecheck(p1) == NumType, "test 1")
assert( typecheck(p2) == TypeError, "test 2")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mPlus[39m([33mConst[39m([32m1[39m), [33mConst[39m([32m2[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mPlus[39m([33mConst[39m([32m1[39m), False))
```

```scala
// test Minus
val p1 = Main( Minus(Const(2), Const(1)) )
val p2 = Main( Minus(Const(2), False) )
assert( typecheck(p1) == NumType, "test 1")
assert( typecheck(p2) == TypeError, "test 2")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mMinus[39m([33mConst[39m([32m2[39m), [33mConst[39m([32m1[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mMinus[39m([33mConst[39m([32m2[39m), False))
```

```scala
// test And
val p1 = Main( And(True, True) )
val p2 = Main( And(Const(2), False) )
assert( typecheck(p1) == BoolType, "test 1")
assert( typecheck(p2) == TypeError, "test 2")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mAnd[39m(True, True))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mAnd[39m([33mConst[39m([32m2[39m), False))
```

```scala
// test Or
val p1 = Main( Or(True, True) )
val p2 = Main( Or(Const(2), False) )
assert( typecheck(p1) == BoolType, "test 1")
assert( typecheck(p2) == TypeError, "test 2")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mOr[39m(True, True))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mOr[39m([33mConst[39m([32m2[39m), False))
```

```scala
// test Neg 
val p1 = Main( Neg(True) )
val p2 = Main( Neg(And(True, True)) )
assert( typecheck(p1) == BoolType, "test 1")
assert( typecheck(p2) == BoolType, "test 2")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mNeg[39m(True))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mNeg[39m([33mAnd[39m(True, True)))
```

```scala
// test Neg
val p1 = Main( Neg(And(True, Or(True, False))) )
val p2 = Main( Neg(And(True, Const(10))) )
val p3 = Main( Neg(Const(1)) )
val p4 = Main( Neg(Minus(Const(1), Const(2))) )
val p5 = Main( Neg(Minus(Const(1), False)) )
assert( typecheck(p1) == BoolType, "test 1")
assert( typecheck(p2) == TypeError, "test 2")
assert( typecheck(p3) == NumType, "test 3")
assert( typecheck(p4) == NumType, "test 4")
assert( typecheck(p5) == TypeError, "test 5")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mNeg[39m([33mAnd[39m(True, [33mOr[39m(True, False))))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mNeg[39m([33mAnd[39m(True, [33mConst[39m([32m10[39m))))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mNeg[39m([33mConst[39m([32m1[39m)))
[36mp4[39m: [32mMain[39m = [33mMain[39m([33mNeg[39m([33mMinus[39m([33mConst[39m([32m1[39m), [33mConst[39m([32m2[39m))))
[36mp5[39m: [32mMain[39m = [33mMain[39m([33mNeg[39m([33mMinus[39m([33mConst[39m([32m1[39m), False)))
```

```scala
// test Geq
val p1 = Main( Geq(Const(1), Const(2)) )
val p2 = Main( Geq(Minus(Const(2), Const(1)), Const(2)) )
val p3 = Main( Geq(And(True, False),  Const(2)) )
val p4 = Main( Geq(Plus(True, Const(3)),  Const(2)) )
val p5 = Main( Geq(Const(2), Plus(True, Const(3))) )
val p6 = Main( Geq(False, False) )
assert( typecheck(p1) == BoolType, "test 1")
assert( typecheck(p2) == BoolType, "test 2")
assert( typecheck(p3) == TypeError, "test 3")
assert( typecheck(p4) == TypeError, "test 4")
assert( typecheck(p5) == TypeError, "test 5")
assert( typecheck(p6) == TypeError, "test 6")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mGeq[39m([33mConst[39m([32m1[39m), [33mConst[39m([32m2[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mGeq[39m([33mMinus[39m([33mConst[39m([32m2[39m), [33mConst[39m([32m1[39m)), [33mConst[39m([32m2[39m)))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mGeq[39m([33mAnd[39m(True, False), [33mConst[39m([32m2[39m)))
[36mp4[39m: [32mMain[39m = [33mMain[39m([33mGeq[39m([33mPlus[39m(True, [33mConst[39m([32m3[39m)), [33mConst[39m([32m2[39m)))
[36mp5[39m: [32mMain[39m = [33mMain[39m([33mGeq[39m([33mConst[39m([32m2[39m), [33mPlus[39m(True, [33mConst[39m([32m3[39m))))
[36mp6[39m: [32mMain[39m = [33mMain[39m([33mGeq[39m(False, False))
```

```scala
// test Eq
val p1 = Main( Eq(Const(1), Const(2)) )
val p2 = Main( Eq(Minus(Const(2), Const(1)), Const(2)) )
val p3 = Main( Eq(And(True, False),  Const(2)) )
val p4 = Main( Eq(Plus(True, Const(3)),  Const(2)) )
val p5 = Main( Eq(Const(2), Plus(True, Const(3))) )
val p6 = Main( Eq(False, False) )
assert( typecheck(p1) == BoolType, "test 1")
assert( typecheck(p2) == BoolType, "test 2")
assert( typecheck(p3) == BoolType, "test 3")
assert( typecheck(p4) == TypeError, "test 4")
assert( typecheck(p5) == TypeError, "test 5")
assert( typecheck(p6) == BoolType, "test 6")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mEq[39m([33mConst[39m([32m1[39m), [33mConst[39m([32m2[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mEq[39m([33mMinus[39m([33mConst[39m([32m2[39m), [33mConst[39m([32m1[39m)), [33mConst[39m([32m2[39m)))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mEq[39m([33mAnd[39m(True, False), [33mConst[39m([32m2[39m)))
[36mp4[39m: [32mMain[39m = [33mMain[39m([33mEq[39m([33mPlus[39m(True, [33mConst[39m([32m3[39m)), [33mConst[39m([32m2[39m)))
[36mp5[39m: [32mMain[39m = [33mMain[39m([33mEq[39m([33mConst[39m([32m2[39m), [33mPlus[39m(True, [33mConst[39m([32m3[39m))))
[36mp6[39m: [32mMain[39m = [33mMain[39m([33mEq[39m(False, False))
```

```scala
// test ITE
val p1 = Main( ITE(False, Const(2), Const(1)) )
val p2 = Main( ITE(False, True, False) )
val p5 = Main( ITE(False, Const(2), False))
assert( typecheck(p1) == NumType, "test 1")
assert( typecheck(p2) == BoolType, "test 2")
assert( typecheck(p5) == TypeError, "test 5")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mITE[39m(False, [33mConst[39m([32m2[39m), [33mConst[39m([32m1[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mITE[39m(False, True, False))
[36mp5[39m: [32mMain[39m = [33mMain[39m([33mITE[39m(False, [33mConst[39m([32m2[39m), False))
```

```scala
// test FunDef
val p1 = Main( FunDef( "x", NumType, Plus(Ident("x"), Const(1))) )
assert( typecheck(p1) == FunType(NumType, NumType), "test 1")
val p2 = Main( FunDef( "x", NumType, Eq(Ident("x"), Const(1))) )
assert( typecheck(p2) == FunType(NumType, BoolType), "test 2")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mFunDef[39m([32m"x"[39m, NumType, [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m))))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mFunDef[39m([32m"x"[39m, NumType, [33mEq[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m))))
```

```scala
// test Let
val p1 = Main( FunDef("g", NumType, Let("h", NumType, Const(1), Mult(Ident("g"), Ident("h")))) )
assert( typecheck(p1) == FunType(NumType, NumType), "test 1")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m(
  [33mFunDef[39m(
    [32m"g"[39m,
    NumType,
    [33mLet[39m([32m"h"[39m, NumType, [33mConst[39m([32m1[39m), [33mMult[39m([33mIdent[39m([32m"g"[39m), [33mIdent[39m([32m"h"[39m)))
  )
)
```

```scala
// test FunCall
val p1 = Main( FunCall(FunDef("a", NumType, Const(1)), Const(1)) )
assert( typecheck(p1) == NumType, "test 1")
val p2 = Main( FunCall(FunDef("b", NumType, Const(1)), Ident("b")) )
assert( typecheck(p2) == TypeError, "test 2")
val p3 = Main( FunCall(FunDef("c", NumType, Ident("d")), Const(1)) )
assert( typecheck(p3) == TypeError, "test 3")
val p4 = Main( FunCall(FunDef("a", NumType, Eq(Ident("a"), Const(1))), Const(1)) )
assert( typecheck(p4) == BoolType, "test 4")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mFunCall[39m([33mFunDef[39m([32m"a"[39m, NumType, [33mConst[39m([32m1[39m)), [33mConst[39m([32m1[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mFunCall[39m([33mFunDef[39m([32m"b"[39m, NumType, [33mConst[39m([32m1[39m)), [33mIdent[39m([32m"b"[39m)))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mFunCall[39m([33mFunDef[39m([32m"c"[39m, NumType, [33mIdent[39m([32m"d"[39m)), [33mConst[39m([32m1[39m)))
[36mp4[39m: [32mMain[39m = [33mMain[39m(
  [33mFunCall[39m([33mFunDef[39m([32m"a"[39m, NumType, [33mEq[39m([33mIdent[39m([32m"a"[39m), [33mConst[39m([32m1[39m))), [33mConst[39m([32m1[39m))
)
```

```scala
// test LetRec
val p1 = Main( LetRec("f", NumType, "x", NumType,
            Plus(FunCall(Ident("f"), Plus(Ident("x"), Const(1))), Const(2)), 
            FunCall(Ident("f"), Const(2))) )
assert( typecheck(p1) == NumType, "test 1")
```

Output:
```
[36mp1[39m: [32mMain[39m = [33mMain[39m(
  [33mLetRec[39m(
    [32m"f"[39m,
    NumType,
    [32m"x"[39m,
    NumType,
    [33mPlus[39m([33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m))), [33mConst[39m([32m2[39m)),
    [33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mConst[39m([32m2[39m))
  )
)
```

<!-- grade_id: cell-31198cd330553090, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-69b68741c9fda826, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-d15f5cb6ed1607da, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-d8ea3308b64e150d, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-83408b1f20330fbc, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-8f34165e14575453, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-bb4cebe2260d7612, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-c4f5dcf0bf8fdca3, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-8037f9727ae33dd4, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-b59243a5817202d2, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-ede2ed89a8f56746, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-90690da2df5d6da5, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-288a52329f6c5428, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-7d0654077dff6d8b, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-d0489819cc1d8319, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-533d45c0493a5630, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-97f81d04838e2546, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-f91107fcf26e2240, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-3ba6c8052c148fff, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-36a0b90b6303d35e, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-89108e7186bf2e7e, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-e0c1267b4358e17e, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-7755f725e7af3e5c, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-a16534250c0ac8f9, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-6089a04546621af9, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-ea1c6a51893fd4d1, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-b3233ea6f1403203, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-0934c6879471b3a3, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-750499f6519c9cb8, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-b9646dcc63b28149, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-3f118b29e49adbec, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-150fee1089cc254b, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-f7fad77d24bd342e, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

```scala

```