```scala

```

# CSCI 3155 Homework 2

Practice with Algebraic Data Types: Natural Numbers and Lists

__Name__:

Qiuyang Fu

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

## Problem 1 (20 Points) : Natural numbers

Natural numbers can be defined in scala using an abstract data type. First we create a type called ``Nat``, and then add two ways of constructing an element of type Nat. A Nat can either be a Zero, or the successor of another already defined Nat. With this type definition, we can construct elements of type Nat.

```scala
sealed trait Nat
case object Zero extends Nat
case class Succ(pred : Nat) extends Nat

val one   = Succ(Zero)
val two   = Succ(one)
val three = Succ(two)
val four  = Succ(three)
val five  = Succ(four)
```

How are two natural numbers added together? The result is obtained by constructing a new natural number that is incremented from zero as many times as both of the operands combined.

```scala
def plus(x : Nat, y : Nat) : Nat = x match {
    case Zero    => y
    case Succ(x) => Succ(plus(x, y))
}

val ans = plus(two,three)
ans == five
```

### Problem 1a (10 points)
Define a function that multiplies two natural numbers:

<!-- grade_id: cell-07775b70f29b720c, points: 0, grade: false, solution: true -->
```scala
// recursive plus, return zero only match zero
def mult(x : Nat, y : Nat) : Nat = 
    x match {
        case Zero    => Zero
        case Succ(p) => plus(mult(p,y), y)
    }
```

<!-- grade_id: cell-6cf23ea68d798237, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( mult(one, four) == four)
assert( mult(two, three) == Succ(five))
assert( mult(Zero, five) == Zero)
passed(10)
//END TEST
```

### Problem 1b (10 points)

Define a function called ``even`` that checks is a Nat is even, returning ``true`` if the number is even.

<!-- grade_id: cell-dcd74ba982b165c7, points: 0, grade: false, solution: true -->
```scala
// statistics list length
def count(l : Nat) : Int = 
    l match {
        case Zero => 0
        case Succ(p) => count(p) + 1
    }
//find the remainder of logarithm 2 and judge whether it is equal to 0
def even(l : Nat) : Boolean = count(l) % 2 == 0
```

<!-- grade_id: cell-c18cb8cff48ffefb, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( even(Zero) == true)
assert( even(one) == false)
assert( even(five) == false)
assert( even(plus(five,five)) == true)
passed(10)
//END TEST
```

## Problem 2 (30 points) : Lists

### Problem 2a (10 points)
Write the definition of list of Nats in scala using using the cons construction we defined in class. If your definition is correct, the lists defined in the test should be created successfully.

<!-- grade_id: cell-05547163ae0a3a81, points: 0, grade: false, solution: true -->
```scala
// define NatList
sealed trait NatList
case object Empty extends NatList
case class Cons(h : Nat, t : NatList) extends NatList
```

<!-- grade_id: cell-c5a1cf21d594aaa1, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
val list1 = Empty
val list2 = Cons(five, Empty)
val list3 = Cons(three, Cons(one, Cons(four, Cons(one, Cons(five, Empty)))))
passed(10)
//END TEST
```

### Problem 2b (5 points)
Write a ``List => Nat`` function called ``length`` that returns the length of a list as a Nat

<!-- grade_id: cell-198204e8fa90413b, points: 0, grade: false, solution: true -->
```scala
// recursive plus, return Zero only match Empty
def length(l : NatList) : Nat = 
    l match {
        case Empty => Zero
        case Cons(h, t) => plus(length(t), one)
    }
```

<!-- grade_id: cell-4851c80fe2d7a6b5, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( length(list1) == Zero )
assert( length(list2) == one )
assert( length(list3) == five )
passed(5)
//END TEST
```

### Problem 2c (5 points)
Write a ``sum`` function which computes the sum of a list and returns a Nat.

<!-- grade_id: cell-5f653fd80422308e, points: 0, grade: false, solution: true -->
```scala
// recursive plus, return Zero only match Empty
def sum(l : NatList) : Nat = 
    l match {
        case Empty => Zero
        case Cons(h, t) => plus(h, sum(t))
    }
```

<!-- grade_id: cell-dd1215a8749c149d, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( sum(list1) == Zero )
assert( sum(list2) == five )
assert( sum(list3) == Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero)))))))))))))) )
passed(5)
//END TEST
```

### Problem 2d (5 points)
Write a function called ``last`` that returns the last element of a list as a singleton list.

[1] => [1]

[1,2,3,4,5] => [5]

What should you do for empty lists?

<!-- grade_id: cell-a5e99da11725afdb, points: 0, grade: false, solution: true -->
```scala
??? // YOUR CODE HERE
```

<!-- grade_id: cell-f3cd3bc0a702d90e, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( last(list1) == Empty )
assert( last(list2) == list2 )
assert( last(list3) == list2 )
passed(5)
//END TEST
```

### Problem 2e (5 points)
Write the ``countdown`` function: given a Nat x, return a list counting down from x to zero.

<!-- grade_id: cell-bcab1cbf421bcac1, points: 0, grade: false, solution: true -->
```scala
// recursive contact, return Cons(Zero, Empty) only match Zero
def countdown(l : Nat) : NatList = 
    l match {
        case Zero => Cons(Zero, Empty)
        case Succ(p) => Cons(Succ(p), countdown(p)) 
    }
```

<!-- grade_id: cell-17ee23e31318b5e4, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert( countdown(five) == 
    Cons( Succ(Succ(Succ(Succ(Succ(Zero))))),
    Cons( Succ(Succ(Succ(Succ(Zero)))),
    Cons( Succ(Succ(Succ(Zero))),
    Cons( Succ(Succ(Zero)), 
    Cons( Succ(Zero),
    Cons(Zero, 
         Empty))))))
)
passed(5)
//END TEST
```

```scala

```