```scala

```

# CSCI 3155 : Assignment 1

__Name__:

Qiuyang Fu

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Homework Instructions

Homework assignments for CSCI 3155 are jupyter notebooks. Each problem in the assignment notebook will have a short description that describes what the problem is asking for, and possibly some code that sets up the problem. There are editable cells labeled ``??? \\YOUR CODE HERE`` where you write your solution. To make sure your solution is correct there are test cells contain tests that your solution is run against.

To complete a problem you first need to run any setup code cells, this executes the code with the scala kernel and adds any created values to the environment so that other cells can refer to them. Then write your solution and execute that cell adding the values created by your solution to the environment. Finally run the test cells. If all tests pass then your code is probably correct, note there may be hidden test that you will be graded against so make sure your code does everything that the problem asks and doesn't just pass the tests. If the test cells don't pass, go back to your solution code and try again, rerunning both the solution cells and test cells to check.

The autograder runs your solutions in order and with a new environment. To make sure that your solution to the assignment works like expected, I suggest restarting the kernel, clearing output, and rerunning all the cells in order. This mimics what the autograder does and helps make sure the code that is tested is what you've written.

Jupyter can sometimes be inconsistent about saving your open notebook to the .ipynb file. To make sure the current version of your assignment is saved, you can use Save and Checkpoint.

Note the autograder expects your homework to be formatted like the assignment handout. Do not duplicate cells, add any cells that your solution uses, delete existing cells, or mess with the format of the assignment. This may cause your assignment to be unable to be graded.

To submit your assignment, upload the .ipynb file with your solutions to the assignment on moodle.

## Problem 1 (25 points) : Installing scala

Install jupyter notebook and almond scala, the scala kernel for jupyter notebook. There are instructions on moodle at: https://moodle.cs.colorado.edu/mod/page/view.php?id=48225.

<!-- grade_id: cell-34b28a7556b0393c, points: 25, grade: true, solution: false -->
```scala
// This is a test cell!

println("Hello World!")
passed(25)
```

Output:
```
Hello World!

*** Tests Passed (25 points) ***
```

## Problem 2 (25 points) : Scala

### Problem 2A (10 points) : Values and variables

Define a thing called ``a`` that is set to ``1`` when it is defined but what it equals can be updated.

Define a thing called ``b`` that is set to ``1``  when it is defined it cannot be changed.

<!-- grade_id: cell-f8cc1d8a3e39bfe5, points: 0, grade: false, solution: true -->
```scala
var a: Int = 1;
val b: Int = 1;
```

Output:
```
[36ma[39m: [32mInt[39m = [32m1[39m
[36mb[39m: [32mInt[39m = [32m1[39m
```

<!-- grade_id: cell-2692da738f666d00, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
a = 2
assert( a == 2)
// This won't compile:
// b = 2
assert( b == 1)
passed(10)
//END TEST
```

Output:
```
*** Tests Passed (10 points) ***
```

### Problem 2B (15 points)

Write a function ``sum`` that computes the sum of the elements of a list.
Write a function ``sumEquals`` that checks whether the sum of a list is is equal to a given number.

<!-- grade_id: cell-998b02508d770320, points: 0, grade: false, solution: true -->
```scala

def sum(x:List[Int]) : Int = {
var s:Int = 0;
for ( a <- x) 
s = s + a;
return s;
}

  
def sumEquals(x:List[Int], s:Int) : Boolean = {
var actual_sum : Int = sum(x); 
if (actual_sum == s) 
    {
return true;
}
return false;
}

```

Output:
```
defined [32mfunction[39m [36msum[39m
defined [32mfunction[39m [36msumEquals[39m
```

<!-- grade_id: cell-f00abb3508bcf67e, points: 15, grade: true, solution: false -->
```scala
//BEGIN TEST
val l1 = List(1,2,3)
val l2 = List()
val l3 = List(1, -2, 3, -4, 5)
assert( sum(l1) == 6)
assert( sum(l2) == 0)
assert( sumEquals(l3, 3))
assert( !sumEquals(l3,2))
passed(15)
//END TEST
```

Output:
```
*** Tests Passed (15 points) ***
[36ml1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m, [32m3[39m)
[36ml2[39m: [32mList[39m[[32mNothing[39m] = [33mList[39m()
[36ml3[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m-2[39m, [32m3[39m, [32m-4[39m, [32m5[39m)
```

```scala

```

```scala

```

```scala

```

```scala

```