```scala

```

# CSCI 3155 : Homework 7



__Name__:

Qiuyang Fu

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1: Functions to Lettuce

### Problem 1a (4 points)

Extend the abstract syntax of lettuce to include functions

$$\begin{array}{rcll}
\mathbf{Expr} & \rightarrow & \\
& ...\\
 & | & FunDef( \mathbf{Identifier}, \mathbf{Expr}) & \texttt{function (arg) body_expr} \\ 
 & | & FunCall(\mathbf{Expr}, \mathbf{Expr}) & \texttt{fun_expr(arg_expr)} \\
\end{array}$$

<!-- grade_id: cell-cf677f18337aa8bf, points: 0, grade: false, solution: true -->
```scala
sealed trait Expr
case class  Const(n : Int) extends Expr
case class  Ident(s : String) extends Expr
case object True extends Expr
case object False extends Expr
case class  Plus( e1 : Expr, e2 : Expr) extends Expr
case class  Minus( e1 : Expr, e2 : Expr) extends Expr
case class  Mult( e1 : Expr, e2 : Expr) extends Expr
case class  Neg( e1 : Expr) extends Expr
case class  And( e1 : Expr, e2 : Expr) extends Expr
case class  Or( e1 : Expr, e2 : Expr) extends Expr
case class  Geq( e1 : Expr, e2 : Expr) extends Expr
case class  Eq( e1 : Expr, e2 : Expr) extends Expr
case class  ITE( ec : Expr, et : Expr, ef : Expr) extends Expr
case class  Let( name : String, ed : Expr, ei : Expr) extends Expr
case class  FunDef( id : String, e : Expr) extends Expr
case class  FunCall( func : Expr, arg : Expr) extends Expr
```

Output:
```
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mIdent[39m
defined [32mobject[39m [36mTrue[39m
defined [32mobject[39m [36mFalse[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mMult[39m
defined [32mclass[39m [36mNeg[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mGeq[39m
defined [32mclass[39m [36mEq[39m
defined [32mclass[39m [36mITE[39m
defined [32mclass[39m [36mLet[39m
defined [32mclass[39m [36mFunDef[39m
defined [32mclass[39m [36mFunCall[39m
```

<!-- grade_id: cell-077f89581a0f7687, points: 4, grade: true, solution: false -->
```scala
val e1 = FunDef( "x", Plus(Ident("x"), Const(1)))
val e2 = FunCall( Ident("f"), Const(1))
val e3 = Let( "f", 
              FunDef( "x", Plus(Ident("x"), Const(1))), 
              FunCall( Ident("f"), Const(1)))
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mFunDef[39m = [33mFunDef[39m([32m"x"[39m, [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m)))
[36me2[39m: [32mFunCall[39m = [33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mConst[39m([32m1[39m))
[36me3[39m: [32mLet[39m = [33mLet[39m(
  [32m"f"[39m,
  [33mFunDef[39m([32m"x"[39m, [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m))),
  [33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mConst[39m([32m1[39m))
)
```

### Problem 1b (4 points)
Recall that our values can be of the form:
- Real numbers: $\mathbb{R}$,
- Booleans: $\mathbb{B}$,
- Closures: $\mathbb{C}$ of the form $\text{Closure}(x, \texttt{e}, \sigma)$,
- Error $\mathbf{error}$.

Extends the defintion of Values to include closures

<!-- grade_id: cell-1119b09fde3fa6bb, points: 0, grade: false, solution: true -->
```scala
type Environment = 
    Map[String, Value]

sealed trait Value
case class NumVal(x : Int) extends Value
case class BinVal(x : Boolean) extends Value
case class Closure(id: String, e: Expr, env: Environment) extends Value
case object Error extends Value
```

Output:
```
defined [32mtype[39m [36mEnvironment[39m
defined [32mtrait[39m [36mValue[39m
defined [32mclass[39m [36mNumVal[39m
defined [32mclass[39m [36mBinVal[39m
defined [32mclass[39m [36mClosure[39m
defined [32mobject[39m [36mError[39m
```

<!-- grade_id: cell-445dc7c6323c566c, points: 4, grade: true, solution: false -->
```scala
val v1 : Value = Closure( "x", Plus(Ident("x"), Const(1)), Map())
val v2 : Value = Closure( "y", Plus(Ident("x"), Ident("y")), Map( "x" -> NumVal(1)))
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36mv1[39m: [32mValue[39m = [33mClosure[39m([32m"x"[39m, [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m)), [33mMap[39m())
[36mv2[39m: [32mValue[39m = [33mClosure[39m([32m"y"[39m, [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)), [33mMap[39m([32m"x"[39m -> [33mNumVal[39m([32m1[39m)))
```

### Problem 1c (20 points)

Extend the eval function to be able to handle non recursive functions.

Recall the inference rules to define the operational semantics of functions in Lettuce:

$$ \begin{array}{c}
 \\
\hline
\sigma \vdash \texttt{FunDef}\ x\ e_b \Downarrow \texttt{Closure}(x, e_b, \sigma) \\
\end{array} \text{(FunDef)} $$

$$ \begin{array}{c}
\sigma \vdash e_f \Downarrow \texttt{Closure}(x, e_b, \pi) \qquad
\sigma \vdash e_a \Downarrow v_a \qquad
\pi[x\rightarrow v_a] \vdash e_b \Downarrow v_r\\
\hline
\sigma \vdash \texttt{FunCall}\ e_f\ e_a \Downarrow v_r \\
\end{array} \text{(FunCall)} $$

<!-- grade_id: cell-e656e0f47dfc3404, points: 0, grade: false, solution: true -->
```scala
def eval( e : Expr, env : Environment) : Value = 
    e match {
        case Const(n)        => NumVal(n) 
        case True            => BinVal(true)
        case False           => BinVal(false)
        case Ident(s)        => if (env.contains(s)) env(s) else Error
        case Plus(e1, e2)    => (eval(e1,env), eval(e2,env)) match {
            case (NumVal(v1), NumVal(v2)) => NumVal( v1 + v2 )
            case _ => Error
        }
        case Minus(e1, e2)   => (eval(e1,env), eval(e2,env)) match {
            case (NumVal(v1), NumVal(v2)) => NumVal( v1 - v2 )
            case _ => Error
        }
        case Mult(e1, e2)    => (eval(e1,env), eval(e2,env)) match {
            case (NumVal(v1), NumVal(v2)) => NumVal( v1 * v2 )
            case _ => Error
        }
        case Neg(e1)         => eval(e1, env) match {
            case NumVal(v) => NumVal(-v)
            case BinVal(v) => BinVal(!v)
            case _ => Error
        }
        case And(e1, e2)     => (eval(e1,env), eval(e2,env)) match {
            case (BinVal(v1), BinVal(v2)) => BinVal( v1 && v2 )
            case _ => Error
        }
        case Or(e1, e2)      => (eval(e1,env), eval(e2,env)) match {
            case (BinVal(v1), BinVal(v2)) => BinVal( v1 || v2 )
            case _ => Error
        }
        case Geq(e1, e2)     => (eval(e1,env), eval(e2,env)) match {
            case (NumVal(v1), NumVal(v2)) => BinVal( v1 >= v2 )
            case _ => Error
        }
        case Eq(e1, e2)      => (eval(e1,env), eval(e2,env)) match {
            case (NumVal(v1), NumVal(v2)) => BinVal( v1 == v2 )
            case (BinVal(v1), BinVal(v2)) => BinVal( v1 == v2 )
            case (BinVal(_), NumVal(_)) => BinVal(false)
            case (NumVal(_), BinVal(_)) => BinVal(false)
            case _ => Error
        }
        case ITE(ec, et, ef) => eval(ec,env) match {
            case BinVal(true) => eval(et,env)
            case BinVal(false) => eval(ef,env)
            case _ => Error
        }
        case Let(name, ed, ei) => eval(ed, env) match {
            case Error => Error
            case vd => {
                val newenv = env + ( name -> vd)
                eval(ei, newenv)
            }
        }
        case FunDef(name, e) => Closure(name, e, env)
        case FunCall(func, arg) => eval(func, env) match {
              case Closure(name, e, env) =>
                val newenv = env + (name -> eval(arg, env))
                eval(e, newenv)
              case BinVal(e) => BinVal(e)
              case NumVal(e) => NumVal(e)
              case Error => Error
        }
    }
```

Output:
```
defined [32mfunction[39m [36meval[39m
```

<!-- grade_id: cell-67c5e0bbe90a30a9, points: 3, grade: true, solution: false -->
```scala
val e1 : Expr = Const(10)
val e2 : Expr = True
val e3 : Expr = Plus( Const(1), Const(2))
val e4 : Expr = And( Geq(e1, e3), e2 )
val e5 : Expr = Let( "x", Const(10), Mult(Ident("x"), Ident("x")))
val emptyenv : Environment = Map()
assert( eval( e1, emptyenv) == NumVal(10) , "test 1")
assert( eval( e2, emptyenv) == BinVal(true), "test 2")
assert( eval( e3, emptyenv) == NumVal(3), "test 3")
assert( eval( e4, emptyenv) == BinVal(true), "test 4")
assert( eval( e5, emptyenv) == NumVal(100), "test 5")
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
[36me1[39m: [32mExpr[39m = [33mConst[39m([32m10[39m)
[36me2[39m: [32mExpr[39m = True
[36me3[39m: [32mExpr[39m = [33mPlus[39m([33mConst[39m([32m1[39m), [33mConst[39m([32m2[39m))
[36me4[39m: [32mExpr[39m = [33mAnd[39m([33mGeq[39m([33mConst[39m([32m10[39m), [33mPlus[39m([33mConst[39m([32m1[39m), [33mConst[39m([32m2[39m))), True)
[36me5[39m: [32mExpr[39m = [33mLet[39m([32m"x"[39m, [33mConst[39m([32m10[39m), [33mMult[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"x"[39m)))
[36memptyenv[39m: [32mEnvironment[39m = [33mMap[39m()
```

<!-- grade_id: cell-95140ee4572efa3c, points: 3, grade: true, solution: false -->
```scala
val e1 : Expr = FunDef("x", Const(10))
val e2 : Expr = FunDef("y", And(Ident("x"), Ident("y")))
val emptyenv : Environment = Map()
assert( eval( e1, emptyenv) == Closure("x", Const(10), emptyenv), "test 6")
assert( eval( e2, emptyenv) == Closure("y", And(Ident("x"), Ident("y")), emptyenv), "test 7")
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
[36me1[39m: [32mExpr[39m = [33mFunDef[39m([32m"x"[39m, [33mConst[39m([32m10[39m))
[36me2[39m: [32mExpr[39m = [33mFunDef[39m([32m"y"[39m, [33mAnd[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)))
[36memptyenv[39m: [32mEnvironment[39m = [33mMap[39m()
```

<!-- grade_id: cell-f766e7537d74199c, points: 3, grade: true, solution: false -->
```scala
val e1 : Expr = FunCall( FunDef("x", Const(10)), True)
val e2 : Expr = FunCall( FunDef("y", And(Ident("x"), Ident("y"))), False)
val emptyenv : Environment = Map()
assert( eval( e1, emptyenv) == NumVal(10), "test 1")
assert( eval( e2, emptyenv) == Error, "test 2")
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
[36me1[39m: [32mExpr[39m = [33mFunCall[39m([33mFunDef[39m([32m"x"[39m, [33mConst[39m([32m10[39m)), True)
[36me2[39m: [32mExpr[39m = [33mFunCall[39m([33mFunDef[39m([32m"y"[39m, [33mAnd[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m))), False)
[36memptyenv[39m: [32mEnvironment[39m = [33mMap[39m()
```

<!-- grade_id: cell-6bfc78c16101e558, points: 3, grade: true, solution: false -->
```scala
val e1 : Expr = FunCall( FunDef("y", Mult( Ident("y"), Const(2))), Const(10))
val e2 : Expr = FunCall( FunDef("z", Neg(Ident("z"))), True)
val e3 : Expr = FunCall( FunDef("z", Neg(Ident("z"))), Const(30))
val emptyenv : Environment = Map()
assert( eval( e1, emptyenv) == NumVal(20), "test 3")
assert( eval( e2, emptyenv) == BinVal(false), "test 2")
assert( eval( e3, emptyenv) == NumVal(-30), "test 3")
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
[36me1[39m: [32mExpr[39m = [33mFunCall[39m([33mFunDef[39m([32m"y"[39m, [33mMult[39m([33mIdent[39m([32m"y"[39m), [33mConst[39m([32m2[39m))), [33mConst[39m([32m10[39m))
[36me2[39m: [32mExpr[39m = [33mFunCall[39m([33mFunDef[39m([32m"z"[39m, [33mNeg[39m([33mIdent[39m([32m"z"[39m))), True)
[36me3[39m: [32mExpr[39m = [33mFunCall[39m([33mFunDef[39m([32m"z"[39m, [33mNeg[39m([33mIdent[39m([32m"z"[39m))), [33mConst[39m([32m30[39m))
[36memptyenv[39m: [32mEnvironment[39m = [33mMap[39m()
```

<!-- grade_id: cell-44ec51e8c2baa94c, points: 4, grade: true, solution: false -->
```scala
val e1 : Expr = Let( "f", 
                     FunDef("x", Plus(Const(1), Ident("x"))),
                     FunCall( Ident("f"), Const(2)))
val e2 : Expr = Let("f",
                   FunDef("x",
                     FunDef("y",
                       Plus(Ident("x"),Ident("y")))),
                   FunCall(FunCall(Ident("f"), Const(2)), Const(3))
                   )
val e3 : Expr = Let("fact", 
                 FunDef("x",
                   ITE( Geq(Const(1), Ident("x")), 
                        Const(1),
                        Mult( Ident("x"),
                              FunCall( Ident("fact"),
                                       Minus(Ident("x"), Const(1)))))),
                 FunCall( Ident("fact"),
                          Const(5))
                   )
val emptyenv : Environment = Map()
assert( eval( e1, emptyenv) == NumVal(3), "test 3")
assert( eval( e2, emptyenv) == NumVal(5), "test 3")
assert( eval( e3, emptyenv) == Error, "test 3")
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mExpr[39m = [33mLet[39m(
  [32m"f"[39m,
  [33mFunDef[39m([32m"x"[39m, [33mPlus[39m([33mConst[39m([32m1[39m), [33mIdent[39m([32m"x"[39m))),
  [33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mConst[39m([32m2[39m))
)
[36me2[39m: [32mExpr[39m = [33mLet[39m(
  [32m"f"[39m,
  [33mFunDef[39m([32m"x"[39m, [33mFunDef[39m([32m"y"[39m, [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)))),
  [33mFunCall[39m([33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mConst[39m([32m2[39m)), [33mConst[39m([32m3[39m))
)
[36me3[39m: [32mExpr[39m = [33mLet[39m(
  [32m"fact"[39m,
  [33mFunDef[39m(
    [32m"x"[39m,
    [33mITE[39m(
      [33mGeq[39m([33mConst[39m([32m1[39m), [33mIdent[39m([32m"x"[39m)),
      [33mConst[39m([32m1[39m),
      [33mMult[39m([33mIdent[39m([32m"x"[39m), [33mFunCall[39m([33mIdent[39m([32m"fact"[39m), [33mMinus[39m
```

<!-- grade_id: cell-728c17e636c94680, points: 4, grade: true, solution: false -->
```scala
val e1 = Let( "y", Const(10),
           Let( "f", FunDef("x", Plus(Ident("x"),Ident("y"))),
             FunCall(Ident("f"), Const(5))))
val e2 = Let( "f", 
           Let( "y", Const(10),
             FunDef("x", Plus(Ident("x"),Ident("y")))),
           FunCall(Ident("f"), Const(5)))
val e3 = Let( "f", 
           Let( "y", Const(10),
             FunDef("x", Plus(Ident("x"),Ident("y")))),
           Let("y", Const(20), 
               FunCall(Ident("f"), Const(5))))
val e4 = Let( "f", 
             FunDef("x", Plus(Ident("x"),Ident("y"))),
           Let("y", Const(20), 
               FunCall(Ident("f"), Const(5))))
val emptyenv : Environment = Map()
assert( eval( e1, emptyenv) == NumVal(15), "test 1")
assert( eval( e2, emptyenv) == NumVal(15), "test 2")
assert( eval( e3, emptyenv) == NumVal(15), "test 2")
assert( eval( e4, emptyenv) == Error, "test 2")
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36me1[39m: [32mLet[39m = [33mLet[39m(
  [32m"y"[39m,
  [33mConst[39m([32m10[39m),
  [33mLet[39m(
    [32m"f"[39m,
    [33mFunDef[39m([32m"x"[39m, [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m))),
    [33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mConst[39m([32m5[39m))
  )
)
[36me2[39m: [32mLet[39m = [33mLet[39m(
  [32m"f"[39m,
  [33mLet[39m([32m"y"[39m, [33mConst[39m([32m10[39m), [33mFunDef[39m([32m"x"[39m, [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)))),
  [33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mConst[39m([32m5[39m))
)
[36me3[39m: [32mLet[39m = [33mLet[39m(
  [32m"f"[39m,
  [33mLet[39m([32m"y"[39m, [33mConst[39m([32m10[39m), [33mFunDef[39m([32m"x"[39m, [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)))),
  [33mLet[39m([32m"y"[39m, [33mConst[39m([32m20[39m), [33mFun
```

## Problem 2 (32 points) : Well Formed Expressions

### Well formed epressions

We have seen many expressions in Lettuce that can Error. Ideally, would like to catch catch errors in Lettuce programs at compile time. One class of errors we have seen are errors related to variable liveness.


In the following programs, there is a variable that is used before it is defined. This causes in error when we try to look up that identifier in the environment.


``x + 1``

``let y = 1 in x + y``

``
let f =
  function(x) x + y
in let y = 
  10
in 
  f(20)``

Variables that are *live* are have been defined before they are used and are in scope when they are used. A *Well Formed* expression is one where all the variables are live. All of the above programs are not well formed.

###  Inference rules to define Well Formed Expressions

Let us define inference rules that will help us check if an expression is checked before use. To do so, we will define a set of currently declared variables that are in scope called $S$. 

We will define inference rules for _judgements_ of the form: $WellFormed(\texttt{e}, S)$ that says that expression `e` is _WellFormed_ under the set of declared in scope variables $S$.

Constants are well formed under any set $S$. 

$$ \begin{array}{c}
\\
\hline
WellFormed(\texttt{Const(f)}, S) \\
\end{array} \text{(const-rule)} $$

The rules for `True` and `False` are similar.

$$ \begin{array}{c}
\\
\hline
WellFormed(\texttt{True}, S) \\
\end{array} \text{(true-rule)} $$

$$ \begin{array}{c}
\\
\hline
WellFormed(\texttt{False}, S) \\
\end{array} \text{(false-rule)} $$

Identifiers are well formed only if they belong to $S$:

$$\begin{array}{c}
x \in S \\
\hline
WellFormed(\texttt{Ident(x)}, S) \\
\end{array} \text{(ident-rule)} $$

The rules for binary operator `Plus, Minus, Mult, Geq, Eq, And, Or` are all very similar.

$$ \begin{array}{c}
WellFormed(\texttt{e1}, S) \;\;\; WellFormed(\texttt{e2}, S)\;\;\; T \in \{ \texttt{Plus}, \texttt{Minus}, \texttt{Mult}, \texttt{Geq}, \texttt{Eq}, \texttt{And}, \texttt{Or} \} \\
\hline
WellFormed(\texttt{T(e1, e2)}, S) \\
\end{array} \text{(binary-op)} $$

Likewise, for unary operators 

$$ \begin{array}{c}
WellFormed(\texttt{e1}, S)\\
\hline
WellFormed(\texttt{Neg(e1)}, S) \\
\end{array} \text{(unary-op)} $$

And the ternary if-then-else operator

$$ \begin{array}{c}
WellFormed(\texttt{e1}, S) \;\;\; WellFormed(\texttt{e2}, S)\;\;\; WellFormed(\texttt{e2}, S)\\
\hline
WellFormed(\texttt{ITE(e1, e2, e3)}, S) \\
\end{array} \text{(ternary-op)} $$

The main rule is for let bindings

$$\begin{array}{c}
WellFormed(\texttt{e1}, S) \;\;\; WellFormed(\texttt{e2}, S \cup \{ x\} ) \\
\hline
WellFormed(\texttt{Let(x, e1, e2)}, S) \\
\end{array} \text{(let-rule)} $$

Suppose `e1` is well formed under $S$ and `e2` is well formed under $S$ with the identifier `x` added,
then `let x = e1 in e2` is well formed under $S$, as well. 
- This is confusing when you first see it, but repeat the sentence above and match it with the rule, 
- Notice that the defining expression must be well formed under the original S but the body expression is allowed to be well formed under $S \cup \{ x\} $.

We can also add the rules for function definition:

$$\begin{array}{c}
WellFormed(\texttt{e}, S \cup \{x\}) \\
\hline
WellFormed(\texttt{FunDef(x, e)}, S ) \\
\end{array} \; \text{(fundef-rule)} $$

Finally, function calls are very simple:

$$ \begin{array}{c}
\texttt{f} \in S,\;\;\;\; WellFormed(\texttt{e}, S) \\
\hline
WellFormed(\texttt{f(e)}, S ) \\
\end{array} \text{(well-formed-funcall)} $$

The name of the function `f` better be defined in the scope $S$ and the argument better be OK. Then
we can say that the function call is.

### Problem 2a

Write a function that determines whether a lettuce expression is well formed based on the inference rules above.

This problem does not have any tests, it will be tested by problem 1b.

<!-- grade_id: cell-0bfb1ef049e169a5, points: 0, grade: false, solution: true -->
```scala
def wellFormed(e: Expr, S: Set[String]): Boolean = e match {
  case Const(_)  => true
  case True => true
  case False => true
  case Ident(x) => S.contains(x)
  case Plus(e1, e2) => wellFormed(e1, S) && wellFormed(e2, S)
  case Minus(e1, e2) => wellFormed(e1, S) && wellFormed(e2, S)
  case Mult(e1, e2) => wellFormed(e1, S) && wellFormed(e2, S)
  case Geq(e1, e2) => wellFormed(e1, S) && wellFormed(e2, S)
  case Eq(e1, e2) => wellFormed(e1, S) && wellFormed(e2, S)
  case And(e1, e2) => wellFormed(e1, S) && wellFormed(e2, S)
  case Or(e1, e2) => wellFormed(e1, S) && wellFormed(e2, S)
  case Neg(e) => wellFormed(e, S)
  case ITE(e1, e2, e3) => wellFormed(e1, S) && wellFormed(e2, S) && wellFormed(e3, S)
  case Let(x, e1, e2) => wellFormed(e1, S) && wellFormed(e2, S + x)
  case FunDef(x, e) => wellFormed(e, S + x)
  case FunCall(f, e) => wellFormed(f, S) && wellFormed(e, S)
}
```

Output:
```
defined [32mfunction[39m [36mwellFormed[39m
```

### Well Formed Lettuce Programs

Recall that a Lettuce program is a lettuce expression that starts executing in an empty environment with no variables declared

$$\mathbf{Program} \rightarrow Main(\mathbf{Expr}) $$

We say that a lettuce program is well formed if its expressions is well formed when no variables are in scope initially. We can make a judgement that says a program `p` is well formed: $WellFormedProgram(\texttt{p})$. And we can define this judgement using inference rules:


$$\begin{array}{c}
WellFormed(\texttt{e}, \emptyset) \\
\hline
WellFormedProgram(\texttt{Main(e)}) \\
\end{array} \text{(wf-program)} $$



```scala
sealed trait Program
case class Main(p : Expr) extends Program
```

Output:
```
defined [32mtrait[39m [36mProgram[39m
defined [32mclass[39m [36mMain[39m
```

### Problem 2b (32 points)

Write a function ``WF`` that determines whether a Lettuce program is well formed based on the inference rules above.

$$ WF : Program \rightarrow Boolean$$

<!-- grade_id: cell-fdf0d5cd0f1a802e, points: 0, grade: false, solution: true -->
```scala
def WF(main: Main) = main match {
  case Main(e: Expr) => wellFormed(e, Set())
  case _ => false
}
```

Output:
```
defined [32mfunction[39m [36mWF[39m
```

<!-- grade_id: cell-a38096a4e3475cb4, points: 2, grade: true, solution: false -->
```scala
val p1 = Main( Const(0) )
val p2 = Main( Const(10) )
val p3 = Main( True )
val p4 = Main( False )
assert( WF(p1), "test 1")
assert( WF(p2), "test 2")
assert( WF(p3), "test 3")
assert( WF(p4), "test 4")
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mConst[39m([32m0[39m))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mConst[39m([32m10[39m))
[36mp3[39m: [32mMain[39m = [33mMain[39m(True)
[36mp4[39m: [32mMain[39m = [33mMain[39m(False)
```

<!-- grade_id: cell-5c48f4cc60a7db27, points: 2, grade: true, solution: false -->
```scala
val p1 = Main( Ident("x") )
val p2 = Main( Ident("y") )
assert( !WF(p1), "test 1")
assert( !WF(p2), "test 2")
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mIdent[39m([32m"x"[39m))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mIdent[39m([32m"y"[39m))
```

<!-- grade_id: cell-fa41127b6d2556f6, points: 1, grade: true, solution: false -->
```scala
val p1 = Main( Plus( Const(10), Const(20)) )
val p2 = Main( Minus( Ident("x"), Ident("x")) )
val p3 = Main( Mult(Plus( Const(10), True), Mult( Const(30), Const(40))) )
val p4 = Main( Mult(Plus( Const(10), Const(20)), Mult( Const(30), Ident("y"))) )
assert(  WF(p1), "test 1")
assert( !WF(p2), "test 2")
assert(  WF(p3), "test 3")
assert( !WF(p4), "test 4")
passed(1)
```

Output:
```
*** Tests Passed (1 point) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mPlus[39m([33mConst[39m([32m10[39m), [33mConst[39m([32m20[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"x"[39m)))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mMult[39m([33mPlus[39m([33mConst[39m([32m10[39m), True), [33mMult[39m([33mConst[39m([32m30[39m), [33mConst[39m([32m40[39m))))
[36mp4[39m: [32mMain[39m = [33mMain[39m([33mMult[39m([33mPlus[39m([33mConst[39m([32m10[39m), [33mConst[39m([32m20[39m)), [33mMult[39m([33mConst[39m([32m30[39m), [33mIdent[39m([32m"y"[39m))))
```

<!-- grade_id: cell-d7ccdae5e5d12006, points: 1, grade: true, solution: false -->
```scala
val p1 = Main( And(And(True, True), And(False,False)) )
val p2 = Main( Or( Or(True, Ident("b1")), Or(Ident("b2"), True)) )
val p3 = Main( And( And( Const(10), True), Or(False, Const(0))) )
assert(  WF(p1), "test 1")
assert( !WF(p2), "test 2")
assert(  WF(p3), "test 3")
passed(1)
```

Output:
```
*** Tests Passed (1 point) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mAnd[39m([33mAnd[39m(True, True), [33mAnd[39m(False, False)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mOr[39m([33mOr[39m(True, [33mIdent[39m([32m"b1"[39m)), [33mOr[39m([33mIdent[39m([32m"b2"[39m), True)))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mAnd[39m([33mAnd[39m([33mConst[39m([32m10[39m), True), [33mOr[39m(False, [33mConst[39m([32m0[39m))))
```

<!-- grade_id: cell-24a7e75d3fd54c53, points: 1, grade: true, solution: false -->
```scala
val p1 = Main( Geq(Const(2), Const(1)))
val p2 = Main( Geq( Plus(Const(1), True), False))
val p3 = Main( Eq( Geq(Const(2), Const(1)), Plus(Const(3), Const(3))))
val p4 = Main( Eq( Geq(Const(2), Const(1)), Plus(Ident("var"), Const(3))))
assert(  WF(p1), "test 1")
assert(  WF(p2), "test 2")
assert(  WF(p3), "test 3")
assert( !WF(p4), "test 4")
passed(1)
```

Output:
```
*** Tests Passed (1 point) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mGeq[39m([33mConst[39m([32m2[39m), [33mConst[39m([32m1[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mGeq[39m([33mPlus[39m([33mConst[39m([32m1[39m), True), False))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mEq[39m([33mGeq[39m([33mConst[39m([32m2[39m), [33mConst[39m([32m1[39m)), [33mPlus[39m([33mConst[39m([32m3[39m), [33mConst[39m([32m3[39m))))
[36mp4[39m: [32mMain[39m = [33mMain[39m([33mEq[39m([33mGeq[39m([33mConst[39m([32m2[39m), [33mConst[39m([32m1[39m)), [33mPlus[39m([33mIdent[39m([32m"var"[39m), [33mConst[39m([32m3[39m))))
```

<!-- grade_id: cell-584caa6466cc7b8e, points: 2, grade: true, solution: false -->
```scala
val p1 = Main( Neg(True) )
val p2 = Main( Neg(Neg(Neg(Neg(Neg(Neg(Const(1))))))) )
val p3 = Main( Neg(Neg(Ident("Neg"))) )
assert(  WF(p1), "test 1")
assert(  WF(p2), "test 2")
assert( !WF(p3), "test 3")
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mNeg[39m(True))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mNeg[39m([33mNeg[39m([33mNeg[39m([33mNeg[39m([33mNeg[39m([33mNeg[39m([33mConst[39m([32m1[39m))))))))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mNeg[39m([33mNeg[39m([33mIdent[39m([32m"Neg"[39m))))
```

<!-- grade_id: cell-6412ecd794d59679, points: 3, grade: true, solution: false -->
```scala
val p1 = Main( ITE(Const(1), Const(2), Const(3)) )
val p2 = Main( ITE(True, Const(1), Ident("f")) )
val p3 = Main( ITE(False, Ident("t"), Const(1)) )
assert(  WF(p1), "test 1")
assert( !WF(p2), "test 2")
assert( !WF(p3), "test 3")
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mITE[39m([33mConst[39m([32m1[39m), [33mConst[39m([32m2[39m), [33mConst[39m([32m3[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mITE[39m(True, [33mConst[39m([32m1[39m), [33mIdent[39m([32m"f"[39m)))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mITE[39m(False, [33mIdent[39m([32m"t"[39m), [33mConst[39m([32m1[39m)))
```

<!-- grade_id: cell-9ba9e789e16b8e24, points: 2, grade: true, solution: false -->
```scala
val p1 = Main(Let( "x", Const(1), Const(1)) )
val p2 = Main(Let( "x", Const(1), Ident("x")) )
val p3 = Main(Let( "x", Plus(Const(1), Const(1)), Plus(Ident("x"), Const(1))) )
val p4 = Main(Let( "x", Ident("x"), Const(1)) )
val p5 = Main(Let( "x", Const(1), Plus( Ident("x"),Ident("y"))) )
assert(  WF(p1), "test 1")
assert(  WF(p2), "test 2")
assert(  WF(p3), "test 3")
assert( !WF(p4), "test 4")
assert( !WF(p5), "test 5")
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mLet[39m([32m"x"[39m, [33mConst[39m([32m1[39m), [33mConst[39m([32m1[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mLet[39m([32m"x"[39m, [33mConst[39m([32m1[39m), [33mIdent[39m([32m"x"[39m)))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mLet[39m([32m"x"[39m, [33mPlus[39m([33mConst[39m([32m1[39m), [33mConst[39m([32m1[39m)), [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m))))
[36mp4[39m: [32mMain[39m = [33mMain[39m([33mLet[39m([32m"x"[39m, [33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m)))
[36mp5[39m: [32mMain[39m = [33mMain[39m([33mLet[39m([32m"x"[39m, [33mConst[39m([32m1[39m), [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m))))
```

<!-- grade_id: cell-7af0a5d3df310b6f, points: 2, grade: true, solution: false -->
```scala
val p1 = Main( Let("x", Const(1), 
                 Let("y", Const(2), 
                   Let("z", Const(3), 
                     Plus( Plus(Ident("x"), Ident("y")), Plus(Ident("z"), Const(4)))))) )
val p2 = Main( Let("x", 
                   Let("y", Const(1), Ident("y")), 
                   Plus(Const(1), Ident("x"))) )
val p3 = Main( Let("x", 
                   Let("y", Const(1), Ident("y")), 
                   Plus(Ident("y"), Ident("x"))) )
assert(  WF(p1), "test 1")
assert(  WF(p2), "test 2")
assert( !WF(p3), "test 3")
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m(
  [33mLet[39m(
    [32m"x"[39m,
    [33mConst[39m([32m1[39m),
    [33mLet[39m(
      [32m"y"[39m,
      [33mConst[39m([32m2[39m),
      [33mLet[39m(
        [32m"z"[39m,
        [33mConst[39m([32m3[39m),
        [33mPlus[39m([33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)), [33mPlus[39m([33mIdent[39m([32m"z"[39m), [33mConst[39m([32m4[39m)))
      )
    )
  )
)
[36mp2[39m: [32mMain[39m = [33mMain[39m(
  [33mLet[39m([32m"x"[39m, [33mLet[39m([32m"y"[39m, [33mConst[39m([32m1[39m), [33mIdent[39m([32m"y"[39m)), [33mPlus[39m([33mConst[39m([32m1[39m), [33mIdent[39m([32m"x"[39m)))
)
[36mp3[39m: [32mMain[39m = [33mMain[39m(
  [33mLet[39m([32m"x"[39m, [33mLet[39m([32m"y"[39m, [33mConst[39m([32m1[39m), [33mIdent[39m([32m"y"[39m)), [33mPlus[39m([33mIdent[39m([32m"y"[39m), [33mIdent[39m([32m"x"
```

<!-- grade_id: cell-f63b788be12ecd2e, points: 4, grade: true, solution: false -->
```scala
val p1 = Main( FunDef("a", Const(1)) )
val p2 = Main( FunDef("b", Ident("b")) )
val p3 = Main( FunDef("c", Ident("d")) )
val p4 = Main( FunDef("e", Mult( Ident("e"), Ident("e"))) )
val p5 = Main( FunDef("f", Eq( Ident("f"), Ident("g"))) )
val p6 = Main( FunDef("g", Let("h", Const(1), Mult(Ident("g"), Ident("h")))) )
assert(  WF(p1), "test 1")
assert(  WF(p2), "test 2")
assert( !WF(p3), "test 3")
assert(  WF(p4), "test 4")
assert( !WF(p5), "test 5")
assert(  WF(p6), "test 6")
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mFunDef[39m([32m"a"[39m, [33mConst[39m([32m1[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mFunDef[39m([32m"b"[39m, [33mIdent[39m([32m"b"[39m)))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mFunDef[39m([32m"c"[39m, [33mIdent[39m([32m"d"[39m)))
[36mp4[39m: [32mMain[39m = [33mMain[39m([33mFunDef[39m([32m"e"[39m, [33mMult[39m([33mIdent[39m([32m"e"[39m), [33mIdent[39m([32m"e"[39m))))
[36mp5[39m: [32mMain[39m = [33mMain[39m([33mFunDef[39m([32m"f"[39m, [33mEq[39m([33mIdent[39m([32m"f"[39m), [33mIdent[39m([32m"g"[39m))))
[36mp6[39m: [32mMain[39m = [33mMain[39m([33mFunDef[39m([32m"g"[39m, [33mLet[39m([32m"h"[39m, [33mConst[39m([32m1[39m), [33mMult[39m([33mIdent[39m([32m"g"[39m), [33mIdent[39m([32m"h"[39m)))))
```

<!-- grade_id: cell-1d993291ec2b99c3, points: 4, grade: true, solution: false -->
```scala
val p1 = Main( FunCall(FunDef("a", Const(1)), Const(1)) )
val p2 = Main( FunCall(FunDef("b", Const(1)), Ident("b")) )
val p3 = Main( FunCall(FunDef("c", Ident("d")), Const(1)) )
assert(  WF(p1), "test 1")
assert( !WF(p2), "test 2")
assert( !WF(p3), "test 3")
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m([33mFunCall[39m([33mFunDef[39m([32m"a"[39m, [33mConst[39m([32m1[39m)), [33mConst[39m([32m1[39m)))
[36mp2[39m: [32mMain[39m = [33mMain[39m([33mFunCall[39m([33mFunDef[39m([32m"b"[39m, [33mConst[39m([32m1[39m)), [33mIdent[39m([32m"b"[39m)))
[36mp3[39m: [32mMain[39m = [33mMain[39m([33mFunCall[39m([33mFunDef[39m([32m"c"[39m, [33mIdent[39m([32m"d"[39m)), [33mConst[39m([32m1[39m)))
```

<!-- grade_id: cell-df1101203b3779b9, points: 4, grade: true, solution: false -->
```scala
val p1 = Main( Let("fact", 
                 FunDef("x",
                   ITE( Geq(Const(1), Ident("x")), 
                        Const(1),
                        Mult( Ident("x"),
                              FunCall( Ident("fact"),
                                       Minus(Ident("x"), Const(1)))))),
                 FunCall( Ident("fact"),
                          Const(5))
                   )
             )
assert( !WF(p1) )
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m(
  [33mLet[39m(
    [32m"fact"[39m,
    [33mFunDef[39m(
      [32m"x"[39m,
      [33mITE[39m(
        [33mGeq[39m([33mConst[39m([32m1[39m), [33mIdent[39m([32m"x"[39m)),
        [33mConst[39m([32m1[39m),
        [33mMult[39m([33mIdent[39m([32m"x"[39m), [33mFunCall[39m([33mIdent[39m([32m"fact"[39m), [33mMinus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m))))
      )
    ),
    [33mFunCall[39m([33mIdent[39m([32m"fact"[39m), [33mConst[39m([32m5[39m))
  )
)
```

<!-- grade_id: cell-58b4311f88f684e2, points: 4, grade: true, solution: false -->
```scala
val p1 = Main( Let("f",
                   FunDef("x",
                     FunDef("y",
                       Plus(Ident("x"),Ident("y")))),
                   FunCall(FunCall(Ident("f"), Const(2)), Const(3))
                   )
              )
assert(  WF(p1) )
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36mp1[39m: [32mMain[39m = [33mMain[39m(
  [33mLet[39m(
    [32m"f"[39m,
    [33mFunDef[39m([32m"x"[39m, [33mFunDef[39m([32m"y"[39m, [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)))),
    [33mFunCall[39m([33mFunCall[39m([33mIdent[39m([32m"f"[39m), [33mConst[39m([32m2[39m)), [33mConst[39m([32m3[39m))
  )
)
```