```scala

```

# CSCI 3155 : Homework 6



__Name__:

Qiuyang Fu

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (20 points): Map and filter with fold


```scala
sealed trait MyList[+A]
case object Empty extends MyList[Nothing]
case class Cons[A]( x : A, xs : MyList[A]) extends MyList[A]

def fold[A,B](f : (B,A) => B, acc : B, list : MyList[A]) : B = 
    list match {
        case Empty => acc
        case Cons(h,t) => fold(f, f(acc,h), t)
    }
```

Output:
```
defined [32mtrait[39m [36mMyList[39m
defined [32mobject[39m [36mEmpty[39m
defined [32mclass[39m [36mCons[39m
defined [32mfunction[39m [36mfold[39m
```

Here are the ``snoc`` and ``concat`` list utility functions. These might be useful for the rest of problem 1.

The function ``snoc`` adds an element to the end of a list : 

``snoc( [1,2,3], 4 ) = [1,2,3,4]``


The function ``concat`` concatenates two lists : 

``concat( [1,2,3], [4,5,6] ) = [1,2,3,4,5,6]``


```scala
def snoc[A]( list : MyList[A], a : A) : MyList[A]= 
    list match {
        case Empty => Cons( a, Empty)
        case Cons( h, t) => Cons( h, snoc(t,a))
    }

snoc( Cons(1, Cons(2, Cons(3, Empty))), 4)
snoc( snoc( snoc( snoc(Empty, 1), 2), 3), 4)

def concat[A]( list1: MyList[A], list2 : MyList[A]) : MyList[A] =
    list2 match {
        case Empty => list1
        case Cons( h, t) => concat( snoc(list1,h), t)
    }

concat( Cons(1, Cons(2, Cons(3, Empty))), Cons(4, Cons(5, Cons(6, Empty))))
```

Output:
```
defined [32mfunction[39m [36msnoc[39m
[36mres2_1[39m: [32mMyList[39m[[32mInt[39m] = [33mCons[39m([32m1[39m, [33mCons[39m([32m2[39m, [33mCons[39m([32m3[39m, [33mCons[39m([32m4[39m, Empty))))
[36mres2_2[39m: [32mMyList[39m[[32mInt[39m] = [33mCons[39m([32m1[39m, [33mCons[39m([32m2[39m, [33mCons[39m([32m3[39m, [33mCons[39m([32m4[39m, Empty))))
defined [32mfunction[39m [36mconcat[39m
[36mres2_4[39m: [32mMyList[39m[[32mInt[39m] = [33mCons[39m(
  [32m1[39m,
  [33mCons[39m([32m2[39m, [33mCons[39m([32m3[39m, [33mCons[39m([32m4[39m, [33mCons[39m([32m5[39m, [33mCons[39m([32m6[39m, Empty)))))
)
```

### Problem 1a

Write the ``map`` function for the ``MyList`` type using ``fold``.

<!-- grade_id: cell-9e11cd3de61ea999, points: 0, grade: false, solution: true -->
```scala
  def map1a[A, B](f: A => B, list: MyList[A]): MyList[B] =
    fold( 
      (left:MyList[B],head:A) => snoc(left,f(head)),
      Empty,
      list
    )
```

Output:
```
defined [32mfunction[39m [36mmap1a[39m
```

<!-- grade_id: cell-515acf15795db239, points: 10, grade: true, solution: false -->
```scala
val list1 = Cons(1, Cons(2, Cons(3, Cons(4, Cons(5, Empty)))))
val list2 = Cons(2, Cons(4, Cons(6, Cons(8, Cons(10, Empty)))))
val list3 = Cons("1", Cons("2", Cons("3", Cons("4", Cons("5", Empty)))))

assert( map1a( (x : Int) => 2 * x, Empty) == Empty , "test 1")
assert( map1a[Int,Int]( x => 2 * x, list1) == list2 , "test 2")
assert( map1a[Int,String]( x => x.toString, list1) == list3, "test 3")
assert( map1a( (x : String) => x.toInt * 2, list3) == list2, "test 4")
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
[36mlist1[39m: [32mCons[39m[[32mInt[39m] = [33mCons[39m([32m1[39m, [33mCons[39m([32m2[39m, [33mCons[39m([32m3[39m, [33mCons[39m([32m4[39m, [33mCons[39m([32m5[39m, Empty)))))
[36mlist2[39m: [32mCons[39m[[32mInt[39m] = [33mCons[39m([32m2[39m, [33mCons[39m([32m4[39m, [33mCons[39m([32m6[39m, [33mCons[39m([32m8[39m, [33mCons[39m([32m10[39m, Empty)))))
[36mlist3[39m: [32mCons[39m[[32mString[39m] = [33mCons[39m(
  [32m"1"[39m,
  [33mCons[39m([32m"2"[39m, [33mCons[39m([32m"3"[39m, [33mCons[39m([32m"4"[39m, [33mCons[39m([32m"5"[39m, Empty))))
)
```

### Problem 1b

Write the ``filter`` function for the ``MyList`` type using ``fold``.

<!-- grade_id: cell-ba1618e5d398ed3d, points: 0, grade: false, solution: true -->
```scala
  def filter1b[A](f: A => Boolean, list: MyList[A]): MyList[A] =
    fold(
       (l : MyList[A], a:A ) => {
          if(f(a)) snoc(l, a) else l
       },
       Empty,
       list
    )
```

Output:
```
defined [32mfunction[39m [36mfilter1b[39m
```

<!-- grade_id: cell-3de5bb68ed294aaa, points: 10, grade: true, solution: false -->
```scala
val list1 = Cons(1, Cons(2, Cons(3, Cons(4, Cons(5, Empty)))))
val list2 = Cons(2, Cons(4, Empty))
val list3 = Cons(1, Cons(3, Cons(5, Empty)))
val list4 = Cons("one", Cons("two", Cons("three", Cons("four", Cons("five", Empty)))))
val list5 = Cons("three", Cons("four", Cons("five", Empty)))

assert( filter1b( (x : Int) => x % 2 == 0, Empty) == Empty, "test 1")
assert( filter1b( (x : Int) => x % 2 == 0, list1) == list2, "test 2")
assert( filter1b[Int]( x => x % 2 == 1, list1) == list3, "test 3")
assert( filter1b( (x : String) => x.length > 3, list4) == list5, "test 4")
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
[36mlist1[39m: [32mCons[39m[[32mInt[39m] = [33mCons[39m([32m1[39m, [33mCons[39m([32m2[39m, [33mCons[39m([32m3[39m, [33mCons[39m([32m4[39m, [33mCons[39m([32m5[39m, Empty)))))
[36mlist2[39m: [32mCons[39m[[32mInt[39m] = [33mCons[39m([32m2[39m, [33mCons[39m([32m4[39m, Empty))
[36mlist3[39m: [32mCons[39m[[32mInt[39m] = [33mCons[39m([32m1[39m, [33mCons[39m([32m3[39m, [33mCons[39m([32m5[39m, Empty)))
[36mlist4[39m: [32mCons[39m[[32mString[39m] = [33mCons[39m(
  [32m"one"[39m,
  [33mCons[39m([32m"two"[39m, [33mCons[39m([32m"three"[39m, [33mCons[39m([32m"four"[39m, [33mCons[39m([32m"five"[39m, Empty))))
)
[36mlist5[39m: [32mCons[39m[[32mString[39m] = [33mCons[39m([32m"three"[39m, [33mCons[39m([32m"four"[39m, [33mCons[39m([32m"five"[39m, Empty)))
```

## Problem 2 (20 points): Map, filter, fold, and lambda practice.

Here are some practice problems using map filter and fold over scala lists. Write the functions as a lambda that takes a list as an argument and then applies one or more of the map filter or fold functions to it to get the result

```scala
// Example 
// get the sum of a list

val sum : List[Int] => Int =
    list =>
        list.foldLeft(0)( (x,y) => x + y)

val test = List(1,2,3,4,5)
assert( sum(test) == 15)
```

Output:
```
[36msum[39m: [32mList[39m[[32mInt[39m] => [32mInt[39m = ammonite.$sess.cmd7$Helper$$Lambda$2524/0x000000080170d840@69dbb51b
[36mtest[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m, [32m3[39m, [32m4[39m, [32m5[39m)
```

<!-- grade_id: cell-8cf280e696124d65, points: 0, grade: false, solution: true -->
```scala
// Problem 2a
// sum all the even numbers in a list

val sumeven : List[Int] => Int = 
    list => 
        list.filter(_ % 2== 0).foldLeft(0)( (x,y) => x + y)
```

Output:
```
[36msumeven[39m: [32mList[39m[[32mInt[39m] => [32mInt[39m = ammonite.$sess.cmd8$Helper$$Lambda$2538/0x0000000801716840@51ba89ac
```

<!-- grade_id: cell-7c23b49a903043df, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1,2,3,4,5)
val test2 = List(108,342,6543,648,217,214)
assert(sumeven(test1) == 6)
assert(sumeven(test2) == 1312)
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mtest1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m, [32m3[39m, [32m4[39m, [32m5[39m)
[36mtest2[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m108[39m, [32m342[39m, [32m6543[39m, [32m648[39m, [32m217[39m, [32m214[39m)
```

<!-- grade_id: cell-03dcf2a103cac464, points: 0, grade: false, solution: true -->
```scala
// Problem 2b
// find the max of a list of positive numbers 

   val max : List[Int] => Int = 
       list =>
        list.foldLeft(0) {
          (x,y) => y.max(x)
        }
```

Output:
```
[36mmax[39m: [32mList[39m[[32mInt[39m] => [32mInt[39m = ammonite.$sess.cmd10$Helper$$Lambda$2564/0x0000000801726040@3a8ee128
```

<!-- grade_id: cell-fc94abed35a74dc1, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1,2,3,4,5)
val test2 = List(345,234,5,76,354,2415,7645,974)
assert(max(test1) == 5)
assert(max(test2) == 7645)
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mtest1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m, [32m3[39m, [32m4[39m, [32m5[39m)
[36mtest2[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m345[39m, [32m234[39m, [32m5[39m, [32m76[39m, [32m354[39m, [32m2415[39m, [32m7645[39m, [32m974[39m)
```

<!-- grade_id: cell-18ce05d1bbf6f67b, points: 0, grade: false, solution: true -->
```scala
// Problem 2c
// replace all odd numbers in a list with their double
val replaceodd : List[Int] => List[Int] = 
    list => list.map { x=>
        if(x % 2 != 0) x * 2 else x
    }
```

Output:
```
[36mreplaceodd[39m: [32mList[39m[[32mInt[39m] => [32mList[39m[[32mInt[39m] = ammonite.$sess.cmd12$Helper$$Lambda$2582/0x0000000801730840@16219ea7
```

<!-- grade_id: cell-0400375448b35757, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1,2,3,4,5)
val test2 = List(17,16,42,25,6,92,18,49)
assert(replaceodd(test1) == List(2,2,6,4,10))
assert(replaceodd(test2) == List(34, 16, 42, 50, 6, 92, 18, 98))
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mtest1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m, [32m3[39m, [32m4[39m, [32m5[39m)
[36mtest2[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m17[39m, [32m16[39m, [32m42[39m, [32m25[39m, [32m6[39m, [32m92[39m, [32m18[39m, [32m49[39m)
```

<!-- grade_id: cell-443681a4843bf4a9, points: 0, grade: false, solution: true -->
```scala
// Problem 2d
// get the multiplicative inverse of a list of numbers, but drop any zeros.

val inverse : List[Double] => List[Double] = 
    list => list.filter( _ != 0.0).map ( 1 / _ )
```

Output:
```
[36minverse[39m: [32mList[39m[[32mDouble[39m] => [32mList[39m[[32mDouble[39m] = ammonite.$sess.cmd14$Helper$$Lambda$2601/0x000000080173c040@19351f58
```

<!-- grade_id: cell-f6e17fcb6c21344b, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1.0,0.0,1.0/3.0,0.0,5.0)
val test2 = List(0.25, 0.0, 4.0)
assert(inverse(test1) == List(1.0, 3.0, 1.0/5))
assert(inverse(test2) == List(4.0, 0.25))
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mtest1[39m: [32mList[39m[[32mDouble[39m] = [33mList[39m([32m1.0[39m, [32m0.0[39m, [32m0.3333333333333333[39m, [32m0.0[39m, [32m5.0[39m)
[36mtest2[39m: [32mList[39m[[32mDouble[39m] = [33mList[39m([32m0.25[39m, [32m0.0[39m, [32m4.0[39m)
```

<!-- grade_id: cell-80158c8add776398, points: 0, grade: false, solution: true -->
```scala
// Problem 2e
// reverse a list

  val reverse : List[Int] => List[Int] = 
        list => 
          list.foldLeft(List[Int]()) {
              (lst,v) => v +:lst
          }
```

Output:
```
[36mreverse[39m: [32mList[39m[[32mInt[39m] => [32mList[39m[[32mInt[39m] = ammonite.$sess.cmd16$Helper$$Lambda$2621/0x0000000801750840@55e1ff05
```

<!-- grade_id: cell-4e71ae8b353b4bef, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1,2,3,4,5)
val test2 = List(36, 29, 71, 61, 21, 66, 47, 27, 98, 60)
assert(reverse(test1) == List(5, 4, 3, 2, 1))
assert(reverse(test2) == List(60, 98, 27, 47, 66, 21, 61, 71, 29, 36))
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mtest1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m, [32m3[39m, [32m4[39m, [32m5[39m)
[36mtest2[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m36[39m, [32m29[39m, [32m71[39m, [32m61[39m, [32m21[39m, [32m66[39m, [32m47[39m, [32m27[39m, [32m98[39m, [32m60[39m)
```

<!-- grade_id: cell-9306547a582cc03c, points: 0, grade: false, solution: true -->
```scala
// Problem 2f
// get the square bracket string representation of a list

  val printlist : List[Int] => String = 
      list => 
          ("[" +: list.foldLeft(List[String]()){
            (l,v) => 
              if (l.isEmpty)  l:+v.toString() else l:+ ", " :+ v.toString 
          } :+ "]").foldLeft("")( _ + _)

```

Output:
```
[36mprintlist[39m: [32mList[39m[[32mInt[39m] => [32mString[39m = ammonite.$sess.cmd18$Helper$$Lambda$2647/0x0000000801760040@6b6cbcf9
```

<!-- grade_id: cell-732b2e0f451e4bff, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1,2,3,4,5)
val test2 = List()
val test3 = List(23, 7, 494, 2884, 17, 6, 38)
assert(printlist(test1) == "[1, 2, 3, 4, 5]")
assert(printlist(test2) == "[]")
assert(printlist(test3) == "[23, 7, 494, 2884, 17, 6, 38]")
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mtest1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m, [32m3[39m, [32m4[39m, [32m5[39m)
[36mtest2[39m: [32mList[39m[[32mNothing[39m] = [33mList[39m()
[36mtest3[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m23[39m, [32m7[39m, [32m494[39m, [32m2884[39m, [32m17[39m, [32m6[39m, [32m38[39m)
```

<!-- grade_id: cell-5345574ae122b1f6, points: 0, grade: false, solution: true -->
```scala
// Problem 2g
// count all points within radius 4 of (0,0)


val circle : List[(Double,Double)] => Int =
      list =>
          list.filter { case (x,y) =>
              (x* x + y * y) <= 16
          }.foldLeft(0) {
            (x ,y ) => x + 1
          }
```

Output:
```
[36mcircle[39m: [32mList[39m[([32mDouble[39m, [32mDouble[39m)] => [32mInt[39m = ammonite.$sess.cmd20$Helper$$Lambda$2674/0x000000080176f840@607756d
```

<!-- grade_id: cell-88066ff345373719, points: 2, grade: true, solution: false -->
```scala
val test1 = List( (0.0,0.0), (1.0,2.0), (3.0,4.0), (5.0,6.0))
val test2 = List( (2.36 , -3.10), (-3.22 , 1.33), (4.23 , 2.99),
                  (1.42 , -3.10), (4.43 , -1.95), (-0.78 , -4.81), 
                  (4.78 , 4.58), (-4.78 , -3.81), (2.62 , 0.04) ,(0.40 , -1.01))
assert(circle(test1) == 2)
assert(circle(test2) == 5)
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mtest1[39m: [32mList[39m[([32mDouble[39m, [32mDouble[39m)] = [33mList[39m(
  ([32m0.0[39m, [32m0.0[39m),
  ([32m1.0[39m, [32m2.0[39m),
  ([32m3.0[39m, [32m4.0[39m),
  ([32m5.0[39m, [32m6.0[39m)
)
[36mtest2[39m: [32mList[39m[([32mDouble[39m, [32mDouble[39m)] = [33mList[39m(
  ([32m2.36[39m, [32m-3.1[39m),
  ([32m-3.22[39m, [32m1.33[39m),
  ([32m4.23[39m, [32m2.99[39m),
  ([32m1.42[39m, [32m-3.1[39m),
  ([32m4.43[39m, [32m-1.95[39m),
  ([32m-0.78[39m, [32m-4.81[39m),
  ([32m4.78[39m, [32m4.58[39m),
  ([32m-4.78[39m, [32m-3.81[39m),
  ([32m2.62[39m, [32m0.04[39m),
  ([32m0.4[39m, [32m-1.01[39m)
)
```

<!-- grade_id: cell-ad40a3f50758f563, points: 0, grade: false, solution: true -->
```scala
// Problem 2h
// calculate calories are in a meal

sealed trait Food
case object Water extends Food
case class IceCream(flavor : String, calories : Int) extends Food
case class Salad(calories : Int) extends Food
case class Sandwich(layers : Int, calories_per_layer : Int) extends Food

  val caloriecount: List[Food] => Int = 
      list =>  
        list.foldLeft(0) {
          case (v, Water) => v
          case (v, IceCream(_,c)) => v + c
          case (v, Salad(c)) => v + c
          case (v, Sandwich(l,pc)) => v + pc * l
        }
```

Output:
```
defined [32mtrait[39m [36mFood[39m
defined [32mobject[39m [36mWater[39m
defined [32mclass[39m [36mIceCream[39m
defined [32mclass[39m [36mSalad[39m
defined [32mclass[39m [36mSandwich[39m
[36mcaloriecount[39m: [32mList[39m[[32mFood[39m] => [32mInt[39m = ammonite.$sess.cmd22$Helper$$Lambda$2720/0x0000000801794840@6b1a715f
```

<!-- grade_id: cell-898c69a568bc7fe4, points: 2, grade: true, solution: false -->
```scala
val meal1 = List( Water, Salad(10))
val meal2 = List( Sandwich(3, 100), IceCream("chocolate", 200), Water, IceCream("vanilla", 300))
assert(caloriecount(meal1) == 10)
assert(caloriecount(meal2) == 800)
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mmeal1[39m: [32mList[39m[[32mProduct[39m with [32mSerializable[39m with [32mFood[39m] = [33mList[39m(Water, [33mSalad[39m([32m10[39m))
[36mmeal2[39m: [32mList[39m[[32mProduct[39m with [32mSerializable[39m with [32mFood[39m] = [33mList[39m(
  [33mSandwich[39m([32m3[39m, [32m100[39m),
  [33mIceCream[39m([32m"chocolate"[39m, [32m200[39m),
  Water,
  [33mIceCream[39m([32m"vanilla"[39m, [32m300[39m)
)
```

<!-- grade_id: cell-5899048198f34809, points: 0, grade: false, solution: true -->
```scala
// Problem 2i
// flatten a list of lists into a single list of elements


  val flatten : List[List[Int]] => List[Int] = 
      lst => lst.foldLeft(List[Int]()) {
        (x , y) => x ++ y
      }
```

Output:
```
[36mflatten[39m: [32mList[39m[[32mList[39m[[32mInt[39m]] => [32mList[39m[[32mInt[39m] = ammonite.$sess.cmd24$Helper$$Lambda$2760/0x00000008017ae840@61a0e4a9
```

<!-- grade_id: cell-836405fa9ec40614, points: 2, grade: true, solution: false -->
```scala
val test1 = List( List(1), List(2,2), List(3,3,3))
val test2 = List( List(43,1), List[Int](), List(9,5,10), List(11))
assert(flatten(test1) == List(1, 2, 2, 3, 3, 3))
assert(flatten(test2) == List(43, 1, 9, 5, 10, 11))
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mtest1[39m: [32mList[39m[[32mList[39m[[32mInt[39m]] = [33mList[39m([33mList[39m([32m1[39m), [33mList[39m([32m2[39m, [32m2[39m), [33mList[39m([32m3[39m, [32m3[39m, [32m3[39m))
[36mtest2[39m: [32mList[39m[[32mList[39m[[32mInt[39m]] = [33mList[39m([33mList[39m([32m43[39m, [32m1[39m), [33mList[39m(), [33mList[39m([32m9[39m, [32m5[39m, [32m10[39m), [33mList[39m([32m11[39m))
```

<!-- grade_id: cell-ca3a805fb51ff2ce, points: 0, grade: false, solution: true -->
```scala
// Problem 2j
// create a map from a list, without adding existing keys

  val makemap : List[(Int,String)] => Map[Int, String] = 
      lst => lst.foldLeft(Map[Int,String]()) {
        case (m ,(k ,v) ) =>
          m.get(k) match {
            case None =>  m + (k->v)
            case Some(_) => m
          }
      }
```

Output:
```
[36mmakemap[39m: [32mList[39m[([32mInt[39m, [32mString[39m)] => [32mMap[39m[[32mInt[39m, [32mString[39m] = ammonite.$sess.cmd26$Helper$$Lambda$2790/0x00000008017c1040@7dae9ca0
```

<!-- grade_id: cell-d863dc73ee02edcb, points: 2, grade: true, solution: false -->
```scala
val test1 = List( (1,"one"), (2, "two"), (3, "three"), (4,"four"), (5, "five"))
val test2 = List( (10, "ten"), (2, "2"), (10, "10"), (2,"two"), (10, "0xA"))
assert(makemap(test1) == Map(1 -> "one", 2 -> "two", 3 -> "three", 4 -> "four",  5 -> "five"))
assert(makemap(test2) == Map(10 -> "ten", 2 -> "2"))
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
[36mtest1[39m: [32mList[39m[([32mInt[39m, [32mString[39m)] = [33mList[39m(
  ([32m1[39m, [32m"one"[39m),
  ([32m2[39m, [32m"two"[39m),
  ([32m3[39m, [32m"three"[39m),
  ([32m4[39m, [32m"four"[39m),
  ([32m5[39m, [32m"five"[39m)
)
[36mtest2[39m: [32mList[39m[([32mInt[39m, [32mString[39m)] = [33mList[39m(
  ([32m10[39m, [32m"ten"[39m),
  ([32m2[39m, [32m"2"[39m),
  ([32m10[39m, [32m"10"[39m),
  ([32m2[39m, [32m"two"[39m),
  ([32m10[39m, [32m"0xA"[39m)
)
```

## Problem 3 (20 points): Dictionaries

We define the datatype $\texttt{Dictionary}$ using the definition below:

$$
\begin{aligned}
\textbf{Dictionary}\ A\ B\ ::&= \text{EmptyDict} \\
&\mid\ \text{Entry}\ A\ B\ (\textbf{Dictionary}\ A\ B)
\end{aligned}
$$

This will create dictionaries of key-value pairs which will be useful for implementing environments. The definition of this type in Scala is given below:

```scala
sealed trait Dictionary[+A, +B] 
case object EmptyDict extends Dictionary[Nothing, Nothing]
case class Entry[A,B](key : A, value : B, xs : Dictionary[A, B]) extends Dictionary[A, B]
```

Output:
```
defined [32mtrait[39m [36mDictionary[39m
defined [32mobject[39m [36mEmptyDict[39m
defined [32mclass[39m [36mEntry[39m
```

### Problem 3a

Write a function called `contains` which returns `True` if the Dictionary contains a given key and `False` if it does not.

$$
\texttt{contains} :  \texttt{Dict}\ a\ b \rightarrow a \rightarrow Boolean
$$

<!-- grade_id: cell-19e9eb529edb079b, points: 0, grade: false, solution: true -->
```scala
def contains[A,B](dict:Dictionary[A,B], key:A) : Boolean = {
    dict match {
        case EmptyDict => false
        case Entry(k, _ , xs) => if (k == key) true else contains(xs,key)
    }
}
```

Output:
```
defined [32mfunction[39m [36mcontains[39m
```

<!-- grade_id: cell-16ed2600c99a7834, points: 5, grade: true, solution: false -->
```scala
assert(contains(Entry(5, true, EmptyDict), 5), "test 1")
assert(contains(Entry(8, true, Entry(5, true, EmptyDict)), 5),"test 2")
assert(!contains(Entry(5, true, EmptyDict), 6), "test 3")
assert(!contains(EmptyDict, 7))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

### Problem 3b

Write a function ``lookup`` that retrieves a value that corresponds to a key in a dictionary. If the key is not in the map then return None.

$$
\texttt{lookup} :  \texttt{Dict}\ a\ b \rightarrow a \rightarrow \texttt{Maybe}\ b
$$

You will need to implement the ``Maybe`` type.

<!-- grade_id: cell-f36e0a24a3d2c697, points: 0, grade: false, solution: true -->
```scala
def lookup[A,B](dict:Dictionary[A,B], key:A):Option[B] = {
    dict match {
        case EmptyDict => None
        case Entry(k, v , xs) => if (k == key ) Some(v) else lookup(xs,key)
    }
}
```

Output:
```
defined [32mfunction[39m [36mlookup[39m
```

<!-- grade_id: cell-ce50b930762797e4, points: 5, grade: true, solution: false -->
```scala
assert(lookup(Entry(5, "ham", EmptyDict), 5) == Some("ham"))
assert(lookup(Entry(8, 7, Entry(5, 3, EmptyDict)), 5) == Some(3))
assert(lookup(Entry(5, true, EmptyDict), 6) == None)
assert(lookup(EmptyDict, 7) == None)
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

### Problem 3d

Create the following dictionary as a scala value using this definition:
```
{
    five: True,
    six: False,
    seven: False
}
```
Place it in a variable named "`my_dict`"

<!-- grade_id: cell-328e9f8e1c4f0c64, points: 0, grade: false, solution: true -->
```scala
val my_dict = 
    Entry(5,true,Entry(6,false,Entry(7,false,EmptyDict)))
```

Output:
```
[36mmy_dict[39m: [32mEntry[39m[[32mInt[39m, [32mBoolean[39m] = [33mEntry[39m(
  [32m5[39m,
  true,
  [33mEntry[39m([32m6[39m, false, [33mEntry[39m([32m7[39m, false, EmptyDict))
)
```

<!-- grade_id: cell-60b14443b5138a7d, points: 2, grade: true, solution: false -->
```scala
assert(lookup(my_dict, 5) == Some(true))
assert(lookup(my_dict, 6) == Some(false))
assert(lookup(my_dict, 7) == Some(false))
assert(lookup(my_dict, 8) == None)
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
```

### Problem 3e
For the following, use the super secret `secret_dict` variable defined below. Use the functions you just defined to complete the following:

1. Look up the value mapped by key `1` and place it in variable `a` 
2. Look up the value mapped by key `3` and place it in variable `b`
3. Look up the value mapped by the key `100` and place it in the variable `c`
4. Check whether the key `-3` is in the map and place True or False in variable `d`
5. Check whether the key `-1` is in the map and place True or False in variable `e`


**Note**: For 1,2,3 remember to extract it from the maybe type, set the variable to the string `"does not exist"` if you get `None` instead of `Some(x)`.

```scala
val secret_dict = Entry(3, "abc",
                        Entry(-1, "xyz",
                              Entry(1, "secret", EmptyDict)))
```

Output:
```
[36msecret_dict[39m: [32mEntry[39m[[32mInt[39m, [32mString[39m] = [33mEntry[39m(
  [32m3[39m,
  [32m"abc"[39m,
  [33mEntry[39m([32m-1[39m, [32m"xyz"[39m, [33mEntry[39m([32m1[39m, [32m"secret"[39m, EmptyDict))
)
```

<!-- grade_id: cell-18d5d9a267c6e266, points: 0, grade: false, solution: true -->
```scala
val a = lookup(secret_dict,1).getOrElse("does not exist")
val b = lookup(secret_dict,3).getOrElse("does not exist")
val c = lookup(secret_dict,100).getOrElse("does not exist")
val d = contains(secret_dict,-3) 
val e = contains(secret_dict,-1) 
```

Output:
```
[36ma[39m: [32mString[39m = [32m"secret"[39m
[36mb[39m: [32mString[39m = [32m"abc"[39m
[36mc[39m: [32mString[39m = [32m"does not exist"[39m
[36md[39m: [32mBoolean[39m = false
[36me[39m: [32mBoolean[39m = true
```

<!-- grade_id: cell-27a2f9fa23da9abb, points: 3, grade: true, solution: false -->
```scala
// Hidden tests (3 pts)
```

### Problem 3f

Write a ``eval`` function for this subset of lettuce using the Dictionaries we have defined to store the environment.


```scala
sealed trait Expr
case class Const(x : Integer) extends Expr
case class Ident(x : String) extends Expr
case class Plus(x : Expr, y : Expr) extends Expr
case class Let(id : String, x : Expr, y : Expr) extends Expr
```

Output:
```
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mLet[39m
```

<!-- grade_id: cell-b0178d64c649d3b3, points: 0, grade: false, solution: true -->
```scala
def eval( expr : Expr, env : Dictionary[String, Int]) : Int = 
    expr match {
        case Const(v) => v
        case Ident(x) => lookup(env,x).get
        case Plus(x,y) => eval(x,env) + eval(y,env)
        case Let(id,x,y) => eval(y,Entry(id,eval(x,env),env))
    }
```

Output:
```
defined [32mfunction[39m [36meval[39m
```

<!-- grade_id: cell-f1840face810867c, points: 5, grade: true, solution: false -->
```scala
val expr1 = Plus( Const(3), Const(4))
val expr2 = Let("x", Const(10),Plus( Ident("x"), Const(1)))
val expr3 = Let("y", Const(20), Let("x", Const(10), Plus( Ident("x"), Ident("y"))))
val expr4 = Let("x", Let("y", Const(10), Plus( Ident("y"), Const(5))), Plus(Ident("x"), Const(5)))

assert(eval(expr1, EmptyDict) == 7)
assert(eval(expr2, EmptyDict) == 11)
assert(eval(expr3, EmptyDict) == 30)
assert(eval(expr4, EmptyDict) == 20)
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36mexpr1[39m: [32mPlus[39m = [33mPlus[39m([33mConst[39m([32m3[39m), [33mConst[39m([32m4[39m))
[36mexpr2[39m: [32mLet[39m = [33mLet[39m([32m"x"[39m, [33mConst[39m([32m10[39m), [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1[39m)))
[36mexpr3[39m: [32mLet[39m = [33mLet[39m(
  [32m"y"[39m,
  [33mConst[39m([32m20[39m),
  [33mLet[39m([32m"x"[39m, [33mConst[39m([32m10[39m), [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)))
)
[36mexpr4[39m: [32mLet[39m = [33mLet[39m(
  [32m"x"[39m,
  [33mLet[39m([32m"y"[39m, [33mConst[39m([32m10[39m), [33mPlus[39m([33mIdent[39m([32m"y"[39m), [33mConst[39m([32m5[39m))),
  [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m5[39m))
)
```

```scala

```