```scala

```

# CSCI 3155 Homework 4

Map, filter and fold.

__Name__:

Qiuyang Fu

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

```scala
sealed trait List[+A]
case object Empty extends List[Nothing]
case class Cons[A]( x : A, xs : List[A]) extends List[A]
```

Output:
```
defined [32mtrait[39m [36mList[39m
defined [32mobject[39m [36mEmpty[39m
defined [32mclass[39m [36mCons[39m
```

## Problem 1 (20 Points) : Map

Write the generic map function. Map takes a $A \rightarrow B$ mapping function and a list of $A$'s and returns the list of $B$'s where the mapping function was applied to each element of the input list.

$$map :  (A \rightarrow B ) \rightarrow List\ A \rightarrow List\ B $$

<!-- grade_id: cell-c0264c2ff9d69e81, points: 0, grade: false, solution: true -->
```scala
def map[A, B](f: (A) => B, l: List[A]) : List[B] = 
    l match {
        case Empty => Empty
        case Cons(x, xs) => Cons(f(x), map[A,B](f, xs))
    }
```

Output:
```
defined [32mfunction[39m [36mmap[39m
```

<!-- grade_id: cell-513c1c34bbfff47c, points: 20, grade: true, solution: false -->
```scala
val list1 = Cons(1, Cons(2, Cons(3, Cons(4, Cons(5, Empty)))))
val list2 = Cons(2, Cons(4, Cons(6, Cons(8, Cons(10, Empty)))))
val list3 = Cons("1", Cons("2", Cons("3", Cons("4", Cons("5", Empty)))))

assert( map( (x : Int) => 2 * x, Empty) == Empty , "test 1")
assert( map[Int,Int]( x => 2 * x, list1) == list2 , "test 2")
assert( map[Int,String]( x => x.toString, list1) == list3, "test 3")
assert( map( (x : String) => x.toInt * 2, list3) == list2, "test 4")
passed(20)
```

Output:
```
*** Tests Passed (20 points) ***
[36mlist1[39m: [32mCons[39m[[32mInt[39m] = [33mCons[39m([32m1[39m, [33mCons[39m([32m2[39m, [33mCons[39m([32m3[39m, [33mCons[39m([32m4[39m, [33mCons[39m([32m5[39m, Empty)))))
[36mlist2[39m: [32mCons[39m[[32mInt[39m] = [33mCons[39m([32m2[39m, [33mCons[39m([32m4[39m, [33mCons[39m([32m6[39m, [33mCons[39m([32m8[39m, [33mCons[39m([32m10[39m, Empty)))))
[36mlist3[39m: [32mCons[39m[[32mString[39m] = [33mCons[39m(
  [32m"1"[39m,
  [33mCons[39m([32m"2"[39m, [33mCons[39m([32m"3"[39m, [33mCons[39m([32m"4"[39m, [33mCons[39m([32m"5"[39m, Empty))))
)
```

## Problem 2 (20 points) : Filter

Write the generic filter function. Filter takes a $A \rightarrow  Boolean$ filtering function and a list of $A$'s and returns a filtered list that only contains the elements that return $true$.

$$ filter : (A \rightarrow \mathbb{B}) \rightarrow List\ A \rightarrow List\ A $$

<!-- grade_id: cell-89fd3b3cc534e8ca, points: 0, grade: false, solution: true -->
```scala
def filter[A](f: (A) => Boolean, l: List[A]) : List[A] = 
    l match {
        case Empty => Empty
        case Cons(x, xs) if (f(x) == true) => Cons(x, filter[A](f, xs))
        case Cons(x, xs) if (f(x) == false) => filter[A](f, xs)
    }
```

Output:
```
defined [32mfunction[39m [36mfilter[39m
```

<!-- grade_id: cell-534162c772c6859a, points: 20, grade: true, solution: false -->
```scala
val list1 = Cons(1, Cons(2, Cons(3, Cons(4, Cons(5, Empty)))))
val list2 = Cons(2, Cons(4, Empty))
val list3 = Cons(1, Cons(3, Cons(5, Empty)))
val list4 = Cons("one", Cons("two", Cons("three", Cons("four", Cons("five", Empty)))))
val list5 = Cons("three", Cons("four", Cons("five", Empty)))

assert( filter( (x : Int) => x % 2 == 0, Empty) == Empty, "test 1")
assert( filter( (x : Int) => x % 2 == 0, list1) == list2, "test 2")
assert( filter[Int]( x => x % 2 == 1, list1) == list3, "test 3")
assert( filter( (x : String) => x.length > 3, list4) == list5, "test 4")
passed(20)
```

Output:
```
*** Tests Passed (20 points) ***
[36mlist1[39m: [32mCons[39m[[32mInt[39m] = [33mCons[39m([32m1[39m, [33mCons[39m([32m2[39m, [33mCons[39m([32m3[39m, [33mCons[39m([32m4[39m, [33mCons[39m([32m5[39m, Empty)))))
[36mlist2[39m: [32mCons[39m[[32mInt[39m] = [33mCons[39m([32m2[39m, [33mCons[39m([32m4[39m, Empty))
[36mlist3[39m: [32mCons[39m[[32mInt[39m] = [33mCons[39m([32m1[39m, [33mCons[39m([32m3[39m, [33mCons[39m([32m5[39m, Empty)))
[36mlist4[39m: [32mCons[39m[[32mString[39m] = [33mCons[39m(
  [32m"one"[39m,
  [33mCons[39m([32m"two"[39m, [33mCons[39m([32m"three"[39m, [33mCons[39m([32m"four"[39m, [33mCons[39m([32m"five"[39m, Empty))))
)
[36mlist5[39m: [32mCons[39m[[32mString[39m] = [33mCons[39m([32m"three"[39m, [33mCons[39m([32m"four"[39m, [33mCons[39m([32m"five"[39m, Empty)))
```

## Problem 3 (20 points) : Fold

Write the generic fold function. Fold takes a an $(B,A) \rightarrow  B$ accumulation function, a type $B$ starting value for the accumulator, and a list of $A$'s. It returns the value of the accumulator after aggregating all the info from the list into the accumulator.

$$ fold : ( (B,A) \rightarrow B ) \rightarrow B \rightarrow List\ A \rightarrow B $$

<!-- grade_id: cell-f61d7dcac19184dc, points: 0, grade: false, solution: true -->
```scala
def fold[A, B](f: (B, A) => B, k: B, l: List[A]) : B = 
    l match {
        case Empty => k
        case Cons(x, Empty) => f(k, x)
        case Cons(x, xs: List[A]) => fold[A, B](f, f(k, x), xs)
    }
```

Output:
```
defined [32mfunction[39m [36mfold[39m
```

<!-- grade_id: cell-5fa7006e32df27db, points: 20, grade: true, solution: false -->
```scala
val list1 = Cons(1, Cons(2, Cons(3, Cons(4, Cons(5, Empty)))))

assert( fold( (acc : Int, x : Int) => acc + x, 15, Empty) == 15,"test 1")
assert( fold( (acc : Int, x : Int) => acc + x, 0, list1) == 15, "test 2")
assert( fold[Int,Int]( (acc, x) => acc - x, 0, list1) == -15, "test 3")
assert( fold( (acc : Int, x : Int) => x - acc, 0, list1) == 3, "test 4")
assert( fold[Int, String]( 
            (acc, x) => if (acc.isEmpty) x.toString else acc + ',' + x.toString, 
            "", 
            list1)
        == "1,2,3,4,5", "test 5")
passed(20)
```

Output:
```
*** Tests Passed (20 points) ***
[36mlist1[39m: [32mCons[39m[[32mInt[39m] = [33mCons[39m([32m1[39m, [33mCons[39m([32m2[39m, [33mCons[39m([32m3[39m, [33mCons[39m([32m4[39m, [33mCons[39m([32m5[39m, Empty)))))
```

```scala

```